Human Brain-Wide Activation of Sleep Rhythms

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Wang et al., recorded concurrent EEG-fMRI in 107 participants during nocturnal NREM sleep to investigate brain activity and connectivity related to slow oscillations (SO), sleep spindles, and in particular their co-occurrence. The authors found SO-spindle coupling to be correlated with increased thalamic and hippocampal activity, and with increased functional connectivity from the hippocampus to the thalamus and from the thalamus to the neocortex, especially the medial prefrontal cortex (mPFC). They concluded the brain-wide activation pattern to resemble episodic memory processing, but to be dissociated from task-related processing and suggest that the thalamus plays a crucial role in coordinating the hippocampal-cortical dialogue during sleep.

The paper offers an impressively large and highly valuable dataset that provides the opportunity for gaining important new insights into the network substrate involved in SOs, spindles, and their coupling.

Thank you for this encouraging assessment. We appreciate your recognition of the value of the dataset and of the questions it allows us to address. Below, we respond to each of your points directly and revise the manuscript accordingly.

Comments on revisions:

Re 1: The revised introduction now cites a couple of papers but discusses them only very superficially, lumping together several studies with very different key results. This is still not very informative for the reader and does not sufficiently acknowledge previously published work. Here are two examples to illustrate this:

(a) "These studies have generally reported that slow oscillations are associated with widespread cortical and subcortical BOLD changes, whereas spindles elicit activation in the thalamus, as well as in several cortical and paralimbic regions." Several studies even showed e.g., a clear activation of the hippocampus and parahippocampal gyrus associated with spindles, not just the thalamus

Thank you for this comment. We agree that our previous sentence was too broad and did not sufficiently reflect the range of findings in the sleep literature. We have therefore rewritten the Introduction to state explicitly that spindle-related BOLD changes have been reported not only in the thalamus, but also in cortical and paralimbic regions, including the hippocampus and parahippocampal gyrus (Bergmann et al., 2012; Caporro et al., 2012; Fogel et al., 2017; Schabus et al., 2007).

Introduction, Page 3-4, Lines 58-62

“Consistent with this view, prior human EEG-fMRI studies have reported spindle-related activation not only in the thalamus, but also in the hippocampus and adjacent parahippocampal gyrus (Bergmann et al., 2012; Schabus et al., 2007). Spindle-related activity has also been linked to striatal engagement, suggesting a broader network that may support memory-related processing during sleep (Fogel et al., 2017).”

Introduction, Page 4, Lines 71-78

“Previous EEG-fMRI studies on sleep have examined both global sleep characteristics (Hale et al., 2016; Moehlman et al., 2019) and the neural correlates of specific waves, including slow oscillations and spindles. These studies have generally shown that slow oscillations are associated with widespread cortical and subcortical BOLD changes (Czisch et al., 2009; Ilhan-Bayrakcı et al., 2022; Picchioni et al., 2011), whereas spindles have been linked not only to thalamic activation but also to cortical and paralimbic regions, including the hippocampus and parahippocampal gyrus (Bergmann et al., 2012; Caporro et al., 2012; Fogel et al., 2017; Schabus et al., 2007).”

Introduction, Page 5, Lines 103-106

“This coupling was associated with increased activation in both the thalamus and hippocampus, with functional connectivity patterns suggesting thalamic coordination of hippocampal-cortical communication, in line with prior EEG-fMRI studies of spindle-related activity (Bergmann et al., 2012; Caporro et al., 2012; Fogel et al., 2017; Schabus et al., 2007).”

(b) "Although these findings provide valuable insights into the BOLD correlates of sleep rhythms, they often do not employ sophisticated temporal modeling (Huang et al., 2024) [, ...]." - previous studies have used e.g., spindle event-related regressors with individual spindle amplitudes as parametric modulators, first and second order derivatives of the HRF function, as well as PPI connectivity analyses, which I would consider rather sophisticated temporal modelling.

We agree that several previous studies have already employed sophisticated modelling approaches, including parametric modulation, HRF derivatives, and PPI analyses (Bergmann et al., 2012; Caporro et al., 2012; Fogel et al., 2017; Picchioni et al., 2011). Our intention was not to suggest that such methods are absent from the literature.

Rather, we aimed to highlight that most prior work has focused on modelling individual SO or spindle events, whereas explicit modelling of their temporal interaction (e.g., SO-spindle coupling) has been less commonly addressed. We have revised the sentence to clarify this point more precisely.

Introduction, Page 4 Lines 78-82

“Although these findings provide important insight into the BOLD correlates of sleep rhythms, most previous studies have focused on individual oscillatory events rather than explicitly modelling their temporal interaction (Bergmann et al., 2012; Caporro et al., 2012; Fogel et al., 2017; Picchioni et al., 2011). Only a few recent studies have begun to examine coupling between rhythms directly, for example Huang et al. (2024).”

Re 4+9: The short overall recordings in some subjects on the one hand and the large number of spindles and SOs detected in N1 sleep stages are still highly concerning, in fact even more so, now that the actual numbers have been provided in the Supplementary Tables. Either the sleep staging or the detection of SO and spindle events must be incorrect. I understand that for specific EEG analysis and fMRI modelling purposes sometimes slightly different thresholds are used as compared to clinical sleep staging, but several parameters here are alarmingly off.

(a) Given that proper NREM sleep (N2+N3) is the relevant stage for the analyses conducted in this paper, some of the N2+N3 durations are very short (eg 7-8 min) while those subjects' results have the same impact on the group level analyses as those with >100 min of N2+N3. Either subjects with very little relevant data (not overall recording time but N2+N3 time) should be excluded or weighting subject data for the group analyses according to the amount od contributed data should be done.

Thank you for the suggestion. It is true that participants with very little N2/3 sleep could contribute noisier subject-level estimates to the group analysis. We therefore checked this directly. Only three participants contributed less than 10 min of N2/3 sleep, and excluding them did not change the main results. For example, hippocampal activation during SO-spindle coupling remained significant after exclusion, t(103) =2.50, p = 0.0071, compared with t(106) = 2.50, p = 0.0070 in the full sample. We have added this control analysis to the Results so that the robustness of the group findings is explicit in the manuscript.

Results, Page 11-12, Lines 238-250

To ensure the results were not driven by individual differences or parameter selection, we conducted a series of control analyses. First, we excluded participants with less than 10 minutes of N2/3 sleep. Only three participants met this criterion, and their exclusion did not change the main results. For example, hippocampal activation during SO-spindle coupling remained significant (t(103) = 2.50, p = 0.0071), comparable to the full sample (t(106) = 2.50, p = 0.0070). Second, because the absolute number of detected SO-spindle coupling events depends on the SO detection threshold, we examined whether the main EEG-fMRI results were sensitive to this parameter. To this end, we varied the SO percentile threshold and reconstructed the EEG-informed GLM at each level. Hippocampal activation during SO-spindle coupling remained significant across a range of thresholds (71st-80th percentile; Fig. S6). Third, to test whether the results depended on the use of a single lateralised frontal electrode, we repeated the EEG-informed fMRI GLM using events detected from Fz. Hippocampal activation during SO-spindle coupling again remained significant (t(106) = 2.47, p = 0.0076), closely matching the original F3-based result (t(106) = 2.50, p = 0.0070).

(b) The authors argue that the SO and spindle detection algorithms are valid since widely used and that they were developed for N2+N3 stages, which is why they will also detect events in other stages: "While, because the detection methods for SO and spindle are based on percentiles, this method will always detect a certain number of events when used for other stages (N1 and REM) sleep data, but the differences between these events and those detected in stage N23 remain unclear." I do agree that with very liberal thresholds, also SO and spindle vents may be detected in other stages, but it shouldn't be that many. If the percentiles of amplitude thresholds were defined based on properly scored N2+N3 stages only, very few events should be detected (erroneously!) in N1, as the occurrence of K-complexes (isolated SOs) and spindles per definition makes it N2, and during REM sleep only very few spindles and SOs are allowed to occur, without scoring it NREM instead. For the first subject (just as example, but with similar numbers for the rest of the sample), reveals as many as 60 SOs and 31 spindles within 8 min of N1 sleep (Table S2) as well as 13 SOs and 7 spindles within 2 min of REM sleep (Table S4). These numbers are completely unrealistic and question the correctness of the sleep staging as well as the physiological relevance of the EEG graphoelements identified as SO and spindles. It also completely undermines the interpretability of the respective event regressors for the fMRI analyses.

(c) Likely, given the large numbers of coupled SO-spindle events and the apparently very low amplitude criteria for event identification, also the number of SO-spindle couplings is likely severely overestimated.

We thank the reviewer for raising this important point. We agree with you that the original stage-wise percentile thresholding could inflate the apparent number of SOs and spindles outside N2/3 sleep. In the original analysis, the thresholds were estimated separately within each sleep stage. As you point out, this procedure can force the detector to label a relatively large number of events in N1 and REM, even when those waveforms should not be interpreted as canonical N2/3 SOs or spindles. We have therefore revised the detection procedure. Following your concern and Reviewer 2’s suggestion, the SO and spindle thresholds are now defined only from N2/3 sleep within each participant, where SOs and spindles are most abundant and physiologically expected to occur. These fixed N2/3-derived thresholds were then applied unchanged to N1 and REM for descriptive reporting. This avoids the artificial normalisation of event detection across sleep stages that can arise when each stage has its own percentile threshold. And we have revised all relevant sections of the manuscript, including “[Results, Page 6-7 Lines 134-148]; [Fig. 1e]; [Results, Page 9 Lines 175-191]; [Fig. 2b]; [Methods, Page 25-27, Lines 567-604]; [Fig. S2-S4]; [Table S2, S4].”

With this revised procedure, detections outside N2/3 are clearly lower than those in N2/3. The mean densities are 2.95 SOs/min, 2.71 spindles/min, and 0.75 coupling events/min in N1, and 2.07 SOs/min, 1.81 spindles/min, and 0.43 coupling events/min in REM. We agree with you that the remaining detections in N1 and REM should not be treated as physiological equivalents of canonical N2/3 SOs, spindles, or SO-spindle complexes. We therefore report them only as descriptive detector outputs obtained under a fixed N2/3-derived threshold (see Table S2, S4 in the revised manuscript). We do not use them to support any physiological claim about SO-spindle coupling in N1 or REM.

This point is also important for the fMRI analyses. You are right that inflated N1 or REM detections would undermine the interpretability of event regressors if those detections entered the EEG-informed fMRI models. They did not. All EEG-informed fMRI GLM and PPI analyses were restricted to N2/3 sleep, where SOs, spindles, and their coupling are physiologically expected and where the detection thresholds were defined. Thus, the central fMRI event regressors were based only on N2/3 events, not on detections from N1 or REM.

We also agree with you that the absolute number of detected SO-spindle couplings depends on the chosen detection threshold. For this reason, we tested whether the main EEG-fMRI result depended on the specific detector setting. Hippocampal activation during SO-spindle coupling remained significant when the SO detection threshold was varied between the 71st and 80th percentiles, as shown in Fig. S6. We therefore do not argue that the detector provides a uniquely correct absolute count of SOs, spindles, or coupling events in every sleep stage. Our conclusion is more specific. The main N2/3 EEG-fMRI finding is robust across a reasonable range of SO detection thresholds, detections in N1 and REM are reported only descriptively, and the physiological interpretation of SO-spindle coupling is restricted to N2/3 sleep.

Results, Page 6-7 Lines 134-148

“Each sleep stage is characterised by distinct spectral properties and rhythmic waveforms, serving as physiological markers (Fig. 1c). Because SO and spindle detection relies on amplitude-based percentile thresholds, we avoided estimating separate thresholds within each sleep stage. Instead, for each participant, the SO and spindle thresholds were defined from N2/3 sleep only, where these rhythms are most abundant and physiologically expected, and the same fixed thresholds were then applied to N1 and REM for descriptive comparison.”

“Under this fixed N2/3-derived thresholding, detected SOs and spindles were larger and more frequent in N2/3 than in N1 or REM. SO and spindle amplitudes were significantly higher during N2/3 sleep (SO: 25.59 ± 1.49 μV; spindle: 7.39 ± 0.27 μV) than during N1 (SO: 20.15 ± 2.32 μV; spindle: 5.23 ± 0.27 μV) and REM sleep (SO: 19.84 ± 1.22 μV; spindle: 5.60 ± 0.22 μV; all p < 1e-4; Fig. 1e, Fig. S2). The corresponding event densities showed the same pattern, with 9.64 ± 0.25 SOs/min and 4.19 ± 0.10 spindles/min in N2/3, compared with 2.95 ± 0.16 SOs/min and 2.71 ± 0.14 spindles/min in N1, and 2.07 ± 0.17 SOs/min and 1.81 ± 0.14 spindles/min in REM (all p < 1e-4). We therefore report detections in N1 and REM only as descriptive outputs of the detector under a fixed N2/3-derived criterion, rather than as physiological equivalents of canonical N2/3 SOs or spindles.”

Fig. 1 legend, Page 8, Line 166-172

“e, Amplitudes (μV) of detected SOs (left) and spindles (right) across sleep stages. SO and spindle detection thresholds were defined from N2/3 sleep within each participant and then applied unchanged to N1 and REM for descriptive comparison. Detections in N1 and REM should therefore be interpreted as detector outputs under this fixed N2/3-derived criterion. The SO amplitudes were measured from the 0.16-1.25 Hz filtered EEG data, and spindle amplitudes were measured from the 12-16 Hz filtered EEG data. Each dot represents an individual participant. Error bars indicate SEM. *** p < 0.001.”

Results, Page 9 Lines 175-191

“SO-spindle coupling is considered important for sleep-dependent memory consolidation. In the current study, using the same N2/3-derived detection thresholds described above, we found that SO-spindle coupling occurred most frequently during N2/3 sleep (2.46 ± 0.06 events/min). Coupling density was significantly lower in N1 (0.75 ± 0.05 events/min, t(106) = 23.54, p < 1e-4) and REM sleep (0.43 ± 0.04 events/min, t(106) = 31.24, p < 1e-4; Fig. 2b, Table S2-S4), consistent with the expected predominance of SO-spindle coupling in NREM sleep (Ngo et al., 2013; Staresina et al., 2015). As with the individual SO and spindle detections, coupling events detected in N1 and REM were retained only for descriptive stage-wise reporting (see Table S2, S4). They were not used to support physiological claims about SO-spindle coupling in these stages, and they were not entered into the EEG-informed fMRI analyses. All subsequent fMRI GLM and PPI analyses were restricted to N2/3 sleep.”

“After extracting all N2/3 EEG epochs in which SO-spindle coupling occurred, we analysed their spectral and phase characteristics. The spindles were most likely to occur slightly before the UP-state peak of SOs (Fig. 2a, e), aligning with results from both animal studies (Maingret et al., 2016) and human research (Staresina et al., 2015). In our data, this pattern was consistent across subjects (Fig. 2d, Rayleigh test: z = 9.51, p < 1e-4), with the peak of the spindle aligned at an SO phase of −41.61 ± 0.86° (the SO UP-state peak is 0°).”

Fig. 2 legend, Page 10, Line 202-205

“b, SO-spindle coupling density across sleep stages, using SO and spindle detections obtained with fixed N2/3-derived thresholds. Coupling events in N1 and REM are shown only for descriptive comparison. The EEG-informed fMRI analyses used N2/3 coupling events only.”

Results, Page 11-12, Lines 242-247

“Second, because the absolute number of detected SO-spindle coupling events depends on the SO detection threshold, we examined whether the main EEG-fMRI results were sensitive to this parameter. To this end, we varied the SO percentile threshold and reconstructed the EEG-informed GLM at each level. Hippocampal activation during SO-spindle coupling remained significant across a range of thresholds (71st - 80th percentile; Fig. S6).”

Methods, Page 25-26, Lines 567-575

“Detection of SOs. Data were first bandpass-filtered between 0.16 and 1.25 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). After identifying all positive-to-negative zero crossings, potential SOs were defined based on the interval between consecutive zero crossings, ranging from 0.8 s to 3 s. For each potential SO, we calculated the amplitude range as the peak minus the trough. For each participant, the amplitude threshold was defined as the 75th percentile of candidate SO amplitude ranges observed during N2/3 sleep. This fixed N2/3-derived threshold was then applied unchanged across the recording for descriptive stage-wise summaries. Detected events were assigned to N1, N2/3 or REM according to the sleep-stage label at the event time. Only candidates exceeding this threshold were labelled as SOs, following previous work (Schreiner et al., 2021).”

Methods, Page 26, Lines 576-583

“Detection of sleep spindles. Detection of sleep spindles. Data were bandpass-filtered between 12 and 16 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). The root mean square (RMS) of the filtered signal was computed with a 200 ms sliding time window. For each participant, the spindle threshold was defined as the 75th percentile of RMS values observed during N2/3 sleep. This fixed N2/3-derived threshold was then applied unchanged across the recording for descriptive stage-wise summaries. Detected events were assigned to N1, N2/3 or REM according to the sleep-stage label at the event time. RMS segments exceeding this threshold for 0.5 s to 3 s were identified as spindles (Staresina et al., 2015).”

Methods, Page 26, Lines 584-591

“Detection of SO-spindle couplings. From the detected SOs and spindles, we identified the peak time of each spindle. Within each SO interval, we checked whether a spindle peak occurred; if so, that SO was labelled as an SO-spindle coupling event. For descriptive stage-wise summaries, coupling events were assigned to the sleep stage of the corresponding SO trough. For every SO-spindle coupling event, an epoch was created time-locked to the SO trough as the central reference, following Schreiner et al. (2021). We extracted data in a [−4 s to 4 s] window around this point, forming the epoch for each coupling event. For the EEG-informed fMRI analyses, only SO, spindle and SO-spindle coupling events detected during N2/3 sleep were used.”

Methods, Page 26-27, Lines 592-604

“The detection procedures described above were developed primarily for N2 and N3 sleep, where SOs, spindles and their coupling are physiologically expected and most reliably observed (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). Because percentile-based thresholds can otherwise force the detector to label events in every sleep stage, we did not estimate separate thresholds within N1 or REM. Instead, for each participant, all SO and spindle thresholds were defined from N2/3 sleep and then applied uniformly across the recording. Tables S1 and S3 report detailed statistical information on sleep rhythm and N2/3 events detection. The N1 and REM events detection reported in Tables S2 and S4, and illustrated in Fig. S2-S4, should therefore be interpreted as descriptive detector outputs under this fixed N2/3-derived criterion, rather than as evidence for canonical N2/3 SOs, spindles or physiological SO-spindle complexes in those stages. These detections were not used in the EEG-informed fMRI GLM or PPI analyses, which were restricted to N2/3 sleep.”

Re 10: The rationale for using a lateralized frontal electrode (F3) for both SO (should have been at least bilateral or central) and spindle detection (should have been a centro-parietal electrode) is not convincing. Other EEG-fMRI spindle or SO papers have used a number of frontal (SO) or centro-parietal (spindles) electrodes averaged or even approaches including all EEG electrodes. Searching events with low thresholds at suboptimal recording sites does not dot this highly valuable dataset justice.

We thank the reviewer for this important comment. We agree that this choice is more sensitive to frontal SOs than to the centro-parietal fast spindle component. Our choice of F3 was driven by the practical constraints of prolonged nocturnal EEG-fMRI recordings. In our MR-compatible EEG setup, FCz was used as the online reference. Central electrodes close to FCz can have reduced signal contrast relative to the reference, and electrodes near the vertex are also more vulnerable to prolonged pressure against the MRI head coil when participants sleep supine for several hours. In this setting, frontal electrodes provided more stable signal quality across the recording. Because the EEG events were used primarily as temporal markers for fMRI modelling, our priority was to obtain reliable event timing during N2/3 sleep rather than to estimate the full scalp topography of SOs and spindles.

We also agree with your concern that this valuable dataset would ideally be analysed with multichannel detection strategies. To test whether the main result depended on the single lateralised F3 site, we repeated the main EEG-informed fMRI analysis using Fz, a midline frontal electrode. The result was unchanged. Hippocampal activation during SO-spindle coupling remained significant when events were detected from Fz, t(106) = 2.47, p = 0.0076, closely matching the original F3-based result, t(106) = 2.50, p = 0.0070. This control analysis does not remove the limitation that centro-parietal fast spindles may be underrepresented, and we do not claim that it does. It does show, however, that the main hippocampal fMRI finding is not driven by idiosyncratic detections from one lateralised frontal electrode. We have made this clearer in the revised manuscript.

Finally, your concern about low thresholds is also important. As described in our response above, the revised analysis now defines SO and spindle thresholds from N2/3 sleep and applies these thresholds uniformly for descriptive comparisons across stages. We also tested the robustness of the hippocampal fMRI result across SO detection thresholds, and the effect remained significant across the 71st to 80th percentile range. We have therefore narrowed the interpretation in the revised manuscript. The main EEG-fMRI result reflects BOLD activity associated with frontal-channel-detected SO-spindle coupling during N2/3 sleep, rather than a full multichannel characterisation of all SO and spindle topographies.

Results, Page 12, Lines 247-250

“Third, to test whether the results depended on the use of a single lateralised frontal electrode, we repeated the EEG-informed fMRI GLM using events detected from Fz. Hippocampal activation during SO-spindle coupling again remained significant (t(106) = 2.47, p = 0.0076), closely matching the original F3-based result (t(106) = 2.50, p = 0.0070).”

Discussion, Page 18, Lines 380-394

“Third, sleep oscillation detection was based on a single frontal electrode. This choice improved signal stability and event timing in the prolonged simultaneous EEG-fMRI setting, but it did not exploit the full multichannel EEG information and cannot characterise the full spatial distribution of SOs and spindles. In particular, F3-based detection may be more sensitive to frontal SOs and frontal sigma activity than to the centro-parietal fast spindle component. We therefore interpret the EEG-informed fMRI results as reflecting BOLD activity associated with frontal-channel-detected SOs, spindles, and their coupling during N2/3 sleep. Future studies using multichannel or source-informed detection strategies, with separate treatment of slow and fast spindles, will be better suited to capture the spatial dynamics of these sleep oscillations. Fourth, the use of large anatomical ROIs may mask subregional contributions of specific thalamic nuclei or hippocampal subfields. Finally, without a memory task, we cannot establish a direct behavioral link between sleep-rhythm-locked activation and memory consolidation. Future studies combining ultra-high-field fMRI or iEEG with cognitive tasks, as well as multichannel or source-informed detection strategies that separately characterize slow and fast spindles, will be better suited to refine our understanding of subregional network dynamics and the functional significance of sleep oscillations.”

Methods, Page 25, Lines 556-566

“It is worth noting that the primary aim of EEG rhythm detection was to identify reliable event times for EEG-informed fMRI modelling. Detection was performed on the F3 electrode because this channel provided stable signal quality during prolonged nocturnal EEG-fMRI recordings. In our MR-compatible EEG setup, FCz was used as the online reference. Central electrodes close to this reference, and electrodes near the vertex that were in prolonged contact with the head coil during supine sleep, were more susceptible to reduced signal contrast, impedance drift, and pressure-related degradation of electrode-scalp contact. We therefore used F3 as a pragmatic choice to maximize reliable event timing in N2/3 sleep. This choice was not intended to characterise the full scalp topography of SOs or spindles, and it may underrepresent the centro-parietal fast spindle component. As a sensitivity analysis, we repeated the main EEG-informed fMRI GLM using Fz, a midline frontal electrode, with the same detection and modelling procedure.”

Re 7: It is not clear to me why/how larger voxels would reduce susceptibility-related distortions and partial volume effects. Usually, the opposite is true. This should be elaborated.

What we meant was that we chose a relatively large voxel size to preserve signal-to-noise ratio and whole-brain coverage within a feasible repetition time for a long overnight EEG-fMRI protocol. This choice is useful for maintaining BOLD sensitivity in sleep recordings, where head motion, physiological noise, and participant comfort are major practical constraints. We agree that it may not be accurate to describe it as reducing susceptibility-related distortion or partial volume effects.

We have rewritten the Methods to state this trade-off directly. The voxel size of 3.5 × 3.5 × 4.2 mm3 allowed whole-brain coverage with a TR of 2000 ms, which was important for modelling sleep-rhythm-related BOLD responses across the whole brain during prolonged nocturnal recordings. A smaller voxel size would have improved spatial specificity, but would also have required either a longer TR, reduced brain coverage, or lower SNR, none of which would have been ideal for the present EEG-fMRI sleep design. We now explicitly acknowledge the cost of this choice.

Methods, Page 21 Lines 453-463

“For the functional scans, whole-brain images were acquired using a T2*-weighted gradient echo-planar imaging (EPI) sequence sensitive to the BOLD contrast. The sequence parameters were as follows: 33 slices in interleaved ascending order, TR = 2000 ms, TE = 30 ms, voxel size = 3.5 × 3.5 × 4.2 mm3, FA = 90°, matrix = 64 × 64, gap = 0.7 mm. A relatively large voxel size was chosen to preserve signal-to-noise ratio while maintaining whole-brain coverage within a feasible repetition time. This compromise was important for the prolonged overnight EEG-fMRI sleep protocol, where head motion, physiological noise, participant comfort, and sustained acquisition stability are substantial practical constraints (Bodurka et al., 2007; Laufs et al., 2008). A smaller voxel size would have improved spatial specificity, but would have required either a longer repetition time, reduced brain coverage, or lower signal-to-noise ratio.”

Reviewer #2 (Public review):

In this study, Wang and colleagues aimed to explore brain-wide activation patterns associated with NREM sleep oscillations, including slow oscillations (SOs), spindles, and SO-spindle coupling events. Their findings reveal that SO-spindle events corresponded with increased activation in both the thalamus and hippocampus. Additionally, they observed that SO-spindle coupling was linked to heightened functional connectivity from the hippocampus to the thalamus, and from the thalamus to the medial prefrontal cortex-three key regions involved in memory consolidation and episodic memory processes.

This study's findings are timely and highly relevant to the field. The authors' extensive data collection, involving 107 participants sleeping in an fMRI while undergoing simultaneous EEG recording, deserves special recognition. If shared, this unique dataset could lead to further valuable insights.

Thank you for this encouraging assessment. We appreciate your recognition of the effort involved in collecting this simultaneous EEG-fMRI sleep dataset. Below, we respond directly to your remaining concern.

Comments on revisions:

The authors' efforts in revising the manuscript and addressing the reviewers' comments are certainly commendable. However, I remain concerned about potential issues in detecting sleep-related oscillations (SOs, spindles, and consequently coupled SO-spindle events), which may arise due to suboptimal parameter selection or inaccurate sleep staging, potentially impacting all subsequent analyses.

A review of Supplementary Tables 1-4 reveals an unusually high number of detected SOs and spindles during sleep stage N1 and REM sleep. While the authors correctly note that a percentile-based detection approach will always identify a certain number of events across sleep stages, the particularly high counts in N1 and REM are concerning. To mitigate the limitations of this method, the authors could have performed event detection independently of sleep stages (i.e., across the entire dataset for each participant) and subsequently assigned the detected events to the corresponding sleep stages. If the event counts in N1 and REM remained disproportionately high, this would indicate a fundamental issue with the detection procedure.

In the previous version, thresholds were estimated separately within each sleep stage. As you point out, this can force the detector to identify a relatively large number of SOs and spindles in N1 and REM, even when those waveforms should not be interpreted as canonical N2/3 events.

We have therefore revised the detection procedure so that event detection is no longer based on separate stage-wise thresholds. Following the logic of your suggestion, we first defined a fixed threshold for each participant and then assigned the detected events to their corresponding sleep stages afterwards. We used N2/3 sleep to define the SO and spindle thresholds because this is the stage in which these events are physiologically expected and most reliably observed. These same N2/3-derived thresholds were then applied unchanged to N1 and REM. This avoids the circularity of forcing a percentile-defined number of detections within each sleep stage.

With this revised procedure, detections outside N2/3 are clearly lower than those in N2/3. The mean densities are 2.95 SOs/min, 2.71 spindles/min, and 0.75 coupling events/min in N1, and 2.07 SOs/min, 1.81 spindles/min, and 0.43 coupling events/min in REM. We also agree with you that the remaining detections in N1 and REM should not be interpreted as physiological equivalents of canonical N2/3 SOs, spindles, or SO-spindle complexes. We now state this explicitly in the manuscript. They are reported only as descriptive detector outputs under the fixed N2/3-derived criterion. And we have revised all relevant sections of the manuscript, including “[Results, Page 6-7 Lines 134-148]; [Fig. 1e]; [Results, Page 9 Lines 175-191]; [Fig. 2b]; [Methods, Page 25-27, Lines 567-604]; [Fig. S2-S4]; [Table S2, S4].”

We also would like to clarify our sleep staging procedure. The sleep staging was first performed using an established automated algorithm, YASA toolkit (Vallat & Walker, 2021), and then manually reviewed by two sleep experts. More importantly for the central results, all EEG-informed fMRI GLM and PPI analyses were restricted to N2/3 sleep. Thus, N1 and REM detections did not enter the event regressors used for the main fMRI analyses and do not affect the interpretation of the hippocampal or thalamic findings.

Finally, we agree that the absolute number of SO-spindle coupling events depends on the detection threshold. We therefore tested whether the main fMRI result depended on the specific SO threshold. Hippocampal activation during SO-spindle coupling remained significant when the SO detection threshold was varied between the 71st and 80th percentiles, as shown in Fig. S6. We have made this clearer in the revised manuscript.

Results, Page 6-7 Lines 134-148

“Each sleep stage is characterised by distinct spectral properties and rhythmic waveforms, serving as physiological markers (Fig. 1c). Because SO and spindle detection relies on amplitude-based percentile thresholds, we avoided estimating separate thresholds within each sleep stage. Instead, for each participant, the SO and spindle thresholds were defined from N2/3 sleep only, where these rhythms are most abundant and physiologically expected, and the same fixed thresholds were then applied to N1 and REM for descriptive comparison.”

“Under this fixed N2/3-derived thresholding, detected SOs and spindles were larger and more frequent in N2/3 than in N1 or REM. SO and spindle amplitudes were significantly higher during N2/3 sleep (SO: 25.59 ± 1.49 μV; spindle: 7.39 ± 0.27 μV) than during N1 (SO: 20.15 ± 2.32 μV; spindle: 5.23 ± 0.27 μV) and REM sleep (SO: 19.84 ± 1.22 μV; spindle: 5.60 ± 0.22 μV; all p < 1e-4; Fig. 1e, Fig. S2). The corresponding event densities showed the same pattern, with 9.64 ± 0.25 SOs/min and 4.19 ± 0.10 spindles/min in N2/3, compared with 2.95 ± 0.16 SOs/min and 2.71 ± 0.14 spindles/min in N1, and 2.07 ± 0.17 SOs/min and 1.81 ± 0.14 spindles/min in REM (all p < 1e-4). We therefore report detections in N1 and REM only as descriptive outputs of the detector under a fixed N2/3-derived criterion, rather than as physiological equivalents of canonical N2/3 SOs or spindles.”

Fig. 1 legend, Page 8, Line 166-172

“e, Amplitudes (μV) of detected SOs (left) and spindles (right) across sleep stages. SO and spindle detection thresholds were defined from N2/3 sleep within each participant and then applied unchanged to N1 and REM for descriptive comparison. Detections in N1 and REM should therefore be interpreted as detector outputs under this fixed N2/3-derived criterion. The SO amplitudes were measured from the 0.16-1.25 Hz filtered EEG data, and spindle amplitudes were measured from the 12-16 Hz filtered EEG data. Each dot represents an individual participant. Error bars indicate SEM. *** p < 0.001.”

Results, Page 9 Lines 175-191

“SO-spindle coupling is considered important for sleep-dependent memory consolidation. In the current study, using the same N2/3-derived detection thresholds described above, we found that SO-spindle coupling occurred most frequently during N2/3 sleep (2.46 ± 0.06 events/min). Coupling density was significantly lower in N1 (0.75 ± 0.05 events/min, t(106) = 23.54, p < 1e-4) and REM sleep (0.43 ± 0.04 events/min, t(106) = 31.24, p < 1e-4; Fig. 2b, Table S2-S4), consistent with the expected predominance of SO-spindle coupling in NREM sleep (Ngo et al., 2013; Staresina et al., 2015). As with the individual SO and spindle detections, coupling events detected in N1 and REM were retained only for descriptive stage-wise reporting (see Table S2, S4). They were not used to support physiological claims about SO-spindle coupling in these stages, and they were not entered into the EEG-informed fMRI analyses. All subsequent fMRI GLM and PPI analyses were restricted to N2/3 sleep.”

“After extracting all N2/3 EEG epochs in which SO-spindle coupling occurred, we analysed their spectral and phase characteristics. The spindles were most likely to occur slightly before the UP-state peak of SOs (Fig. 2a, e), aligning with results from both animal studies (Maingret et al., 2016) and human research (Staresina et al., 2015). In our data, this pattern was consistent across subjects (Fig. 2d, Rayleigh test: z = 9.51, p < 1e-4), with the peak of the spindle aligned at an SO phase of −41.61 ± 0.86° (the SO UP-state peak is 0°).”

Fig. 2 legend, Page 10, Line 202-205

“b, SO-spindle coupling density across sleep stages, using SO and spindle detections obtained with fixed N2/3-derived thresholds. Coupling events in N1 and REM are shown only for descriptive comparison. The EEG-informed fMRI analyses used N2/3 coupling events only.”

Results, Page 11-12, Lines 242-247

“Second, because the absolute number of detected SO-spindle coupling events depends on the SO detection threshold, we examined whether the main EEG-fMRI results were sensitive to this parameter. To this end, we varied the SO percentile threshold and reconstructed the EEG-informed GLM at each level. Hippocampal activation during SO-spindle coupling remained significant across a range of thresholds (71st - 80th percentile; Fig. S6).”

Methods, Page 25-26, Lines 567-575

“Detection of SOs. Data were first bandpass-filtered between 0.16 and 1.25 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). After identifying all positive-to-negative zero crossings, potential SOs were defined based on the interval between consecutive zero crossings, ranging from 0.8 s to 3 s. For each potential SO, we calculated the amplitude range as the peak minus the trough. For each participant, the amplitude threshold was defined as the 75th percentile of candidate SO amplitude ranges observed during N2/3 sleep. This fixed N2/3-derived threshold was then applied unchanged across the recording for descriptive stage-wise summaries. Detected events were assigned to N1, N2/3 or REM according to the sleep-stage label at the event time. Only candidates exceeding this threshold were labelled as SOs, following previous work (Schreiner et al., 2021).”

Methods, Page 26, Lines 576-583

“Detection of sleep spindles. Detection of sleep spindles. Data were bandpass-filtered between 12 and 16 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). The root mean square (RMS) of the filtered signal was computed with a 200 ms sliding time window. For each participant, the spindle threshold was defined as the 75th percentile of RMS values observed during N2/3 sleep. This fixed N2/3-derived threshold was then applied unchanged across the recording for descriptive stage-wise summaries. Detected events were assigned to N1, N2/3 or REM according to the sleep-stage label at the event time. RMS segments exceeding this threshold for 0.5 s to 3 s were identified as spindles (Staresina et al., 2015).”

Methods, Page 26, Lines 584-591

“Detection of SO-spindle couplings. From the detected SOs and spindles, we identified the peak time of each spindle. Within each SO interval, we checked whether a spindle peak occurred; if so, that SO was labelled as an SO-spindle coupling event. For descriptive stage-wise summaries, coupling events were assigned to the sleep stage of the corresponding SO trough. For every SO-spindle coupling event, an epoch was created time-locked to the SO trough as the central reference, following Schreiner et al. (2021). We extracted data in a [−4 s to 4 s] window around this point, forming the epoch for each coupling event. For the EEG-informed fMRI analyses, only SO, spindle and SO-spindle coupling events detected during N2/3 sleep were used.”

Methods, Page 26-27, Lines 592-604

“The detection procedures described above were developed primarily for N2 and N3 sleep, where SOs, spindles and their coupling are physiologically expected and most reliably observed (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). Because percentile-based thresholds can otherwise force the detector to label events in every sleep stage, we did not estimate separate thresholds within N1 or REM. Instead, for each participant, all SO and spindle thresholds were defined from N2/3 sleep and then applied uniformly across the recording. Tables S1 and S3 report detailed statistical information on sleep rhythm and N2/3 events detection. The N1 and REM events detection reported in Tables S2 and S4, and illustrated in Fig. S2-S4, should therefore be interpreted as descriptive detector outputs under this fixed N2/3-derived criterion, rather than as evidence for canonical N2/3 SOs, spindles or physiological SO-spindle complexes in those stages. These detections were not used in the EEG-informed fMRI GLM or PPI analyses, which were restricted to N2/3 sleep.”

Reviewer #3 (Public review):

Summary:

Wang et al., examined the brain activity patterns during sleep, especially when locked to those canonical sleep rhythms such as SO, spindle, and their coupling. Analyzing data from a large sample, the authors found significant coupling between spindles and SOs, particularly during the up-state of the SO. Moreover, the authors examined the patterns of whole-brain activity locked to these sleep rhythms. The authors next investigated the functional connectivity analyses, and found enhanced connectivity between the hippocampus and the thalamus and the medial PFC. These results reinforced the theoretical model of sleep-dependent memory consolidation, such that SO-spindle coupling is conducive for systems-level memory reactivation and consolidation.

Strengths:

There are obvious strengths in this work, including the large sample size, state-of-the-art neuroimaging and neural oscillation analyses, and the richness of results. The results now inform hemodynamic neural activity that coincided with SO-spindle couplings.

Weaknesses:

My earlier comments were about the inability to make inferences on memory given the lack of memory tasks, and the weakness in using the open-ended cognitive state decoding.

Comments on revisions:

The current revision has addressed these major concerns. The authors expanded discussions regarding the theoretical implications of the work in a more nuanced manner.

Thank you for taking the time to re-evaluate the manuscript. We are pleased that the revised Discussion now reads as more nuanced, especially in relation to the limits of the memory-related interpretation. Your earlier comments helped us sharpen both the claims and the framing, and we are grateful for that.

References:

Bergmann, T. O., Mölle, M., Diedrichs, J., Born, J., & Siebner, H. R. (2012). Sleep spindle-related reactivation of category-specific cortical regions after learning face-scene associations. Neuroimage, 59(3), 2733-2742.

Bodurka, J., Ye, F., Petridou, N., Murphy, K., & Bandettini, P. A. (2007). Mapping the MRI voxel volume in which thermal noise matches physiological noise—implications for fMRI. Neuroimage, 34(2), 542-549.

Caporro, M., Haneef, Z., Yeh, H. J., Lenartowicz, A., Buttinelli, C., Parvizi, J., & Stern, J. M. (2012). Functional MRI of sleep spindles and K-complexes. Clinical neurophysiology, 123(2), 303-309.

Czisch, M., Wehrle, R., Stiegler, A., Peters, H., Andrade, K., Holsboer, F., & Sämann, P. G. (2009). Acoustic oddball during NREM sleep: a combined EEG/fMRI study. PloS one, 4(8), e6749.

Fogel, S., Albouy, G., King, B. R., Lungu, O., Vien, C., Bore, A., Pinsard, B., Benali, H., Carrier, J., & Doyon, J. (2017). Reactivation or transformation? Motor memory consolidation associated with cerebral activation time-locked to sleep spindles. PloS one, 12(4), e0174755.

Hahn, M. A., Heib, D., Schabus, M., Hoedlmoser, K., & Helfrich, R. F. (2020). Slow oscillation-spindle coupling predicts enhanced memory formation from childhood to adolescence. Elife, 9, e53730.

Hale, J. R., White, T. P., Mayhew, S. D., Wilson, R. S., Rollings, D. T., Khalsa, S., Arvanitis, T. N., & Bagshaw, A. P. (2016). Altered thalamocortical and intra-thalamic functional connectivity during light sleep compared with wake. Neuroimage, 125, 657-667.

Helfrich, R. F., Lendner, J. D., Mander, B. A., Guillen, H., Paff, M., Mnatsakanyan, L., Vadera, S., Walker, M. P., Lin, J. J., & Knight, R. T. (2019). Bidirectional prefrontal-hippocampal dynamics organize information transfer during sleep in humans. Nature Communications, 10(1), 3572.

Helfrich, R. F., Mander, B. A., Jagust, W. J., Knight, R. T., & Walker, M. P. (2018). Old brains come uncoupled in sleep: slow wave-spindle synchrony, brain atrophy, and forgetting. Neuron, 97(1), 221-230. e224.

Huang, Q., Xiao, Z., Yu, Q., Luo, Y., Xu, J., Qu, Y., Dolan, R., Behrens, T., & Liu, Y. (2024). Replay-triggered brain-wide activation in humans. Nature Communications, 15(1), 7185.

Ilhan-Bayrakcı, M., Cabral-Calderin, Y., Bergmann, T. O., Tüscher, O., & Stroh, A. (2022). Individual slow wave events give rise to macroscopic fMRI signatures and drive the strength of the BOLD signal in human resting-state EEG-fMRI recordings. Cerebral Cortex, 32(21), 4782-4796.

Laufs, H., Daunizeau, J., Carmichael, D. W., & Kleinschmidt, A. (2008). Recent advances in recording electrophysiological data simultaneously with magnetic resonance imaging. Neuroimage, 40(2), 515-528.

Maingret, N., Girardeau, G., Todorova, R., Goutierre, M., & Zugaro, M. (2016). Hippocampo-cortical coupling mediates memory consolidation during sleep. Nature Neuroscience, 19(7), 959-964.

Moehlman, T. M., de Zwart, J. A., Chappel-Farley, M. G., Liu, X., McClain, I. B., Chang, C., Mandelkow, H., Özbay, P. S., Johnson, N. L., & Bieber, R. E. (2019). All-night functional magnetic resonance imaging sleep studies. Journal of neuroscience methods, 316, 83-98.

Ngo, H.-V., Fell, J., & Staresina, B. (2020). Sleep spindles mediate hippocampal-neocortical coupling during long-duration ripples. Elife, 9, e57011.

Ngo, H. V., Martinetz, T., Born, J., & Molle, M. (2013). Auditory closed-loop stimulation of the sleep slow oscillation enhances memory. Neuron, 78(3), 545-553.

Picchioni, D., Horovitz, S. G., Fukunaga, M., Carr, W. S., Meltzer, J. A., Balkin, T. J., Duyn, J. H., & Braun, A. R. (2011). Infraslow EEG oscillations organize large-scale cortical–subcortical interactions during sleep: a combined EEG/fMRI study. Brain research, 1374, 63-72.

Schabus, M., Dang-Vu, T. T., Albouy, G., Balteau, E., Boly, M., Carrier, J., Darsaud, A., Degueldre, C., Desseilles, M., & Gais, S. (2007). Hemodynamic cerebral correlates of sleep spindles during human non-rapid eye movement sleep. Proceedings of the National Academy of Sciences, 104(32), 13164-13169.

Schreiner, T., Kaufmann, E., Noachtar, S., Mehrkens, J.-H., & Staudigl, T. (2022). The human thalamus orchestrates neocortical oscillations during NREM sleep. Nature Communications, 13(1), 5231.

Schreiner, T., Petzka, M., Staudigl, T., & Staresina, B. P. (2021). Endogenous memory reactivation during sleep in humans is clocked by slow oscillation-spindle complexes. Nature Communications, 12(1), 3112.

Staresina, B. P., Bergmann, T. O., Bonnefond, M., van der Meij, R., Jensen, O., Deuker, L., Elger, C. E., Axmacher, N., & Fell, J. (2015). Hierarchical nesting of slow oscillations, spindles and ripples in the human hippocampus during sleep. Nature Neuroscience, 18(11), 1679-1686.

Staresina, B. P., Niediek, J., Borger, V., Surges, R., & Mormann, F. (2023). How coupled slow oscillations, spindles and ripples coordinate neuronal processing and communication during human sleep. Nature Neuroscience, 1-9.

Vallat, R., & Walker, M. P. (2021). An open-source, high-performance tool for automated sleep staging. Elife, 10.

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

During sleep, our brain undergoes highly synchronized activity, orchestrated by distinct neural rhythms. Little is known about the associated brain activation during these sleep rhythms, and even less about their functional implications. In this study, we investigated the brain-wide activation underlying human sleep rhythms by employing simultaneous electroencephalography (EEG) and functional magnetic resonance imaging (fMRI) in 107 participants during nocturnal nap (first half of the night). We identified robust coupling between slow oscillations (SOs) and fast spindles during deep non-rapid eye movement (NREM) sleep (N2/3 stages), with spindle peaks consistently occurring just before the SO UP-state. This SO-spindle coupling was linked to elevated activation in both the thalamus and hippocampus, alongside increased functional connectivity from the hippocampus to the thalamus and from the thalamus to the medial prefrontal cortex (mPFC). An open-ended cognitive state decoding analysis suggested that these activations may relate to episodic memory processes, yet were distinct from task-related networks. Together, these findings highlight the thalamus as a key coordinator of hippocampal-cortical communication during sleep and provide new insights into the mechanisms by which synchronized sleep rhythms may support memory consolidation.

Introduction

During sleep, the human brain engages in highly synchronized activity driven by distinct neural rhythms. In non-rapid eye movement (NREM) sleep, the emergence and coordination of slow oscillations (SOs) and sleep spindles are key features (Hahn et al., 2020; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Staresina et al., 2015; Staresina et al., 2023). However, the brain activation patterns associated with these rhythms, and especially their coordination, remain largely unexplored, and their functional implications are even less understood.

Earlier version: SOs (∼1 Hz) are characterized by fluctuations in neuronal membrane potential between periods of silence (’hyperpolarization’, or ‘DOWN’ state) and excitation (’depolarization’, or ‘UP’ state) (Amzica & Steriade, 2002; Steriade, McCormick, & Sejnowski, 1993). These oscillations originate in neocortical regions, particularly the medial prefrontal cortex (mPFC), during NREM sleep (Dang-Vu et al., 2008; Massimini et al., 2004; Nir et al., 2011). The depolarization phase of SOs is believed to trigger sleep spindles in the thalamus, producing 11–16 Hz oscillations (Fernandez & Lüthi, 2020; Mak-McCully et al., 2017). Thalamic spindles tend to align with the excitable UP-states of SOs and can synchronize with hippocampal ripples (Buzsáki, 2015; Helfrich et al., 2019; Joo & Frank, 2018; Ngo, Fell, & Staresina, 2020). Rodent optogenetics studies have shown that only thalamic spindles phase-locked to the UP-state of cortical SOs enhance memory consolidation, while out-of-phase spindles do not (Latchoumane et al., 2017). This finding is consistent with intracranial evidence from human epilepsy patients, suggesting the thalamus acts as a key relay between the hippocampus and cortex, coordinating their interaction (Coulon, Budde, & Pape, 2012; Ferraris et al., 2021; Schreiner et al., 2022; Staresina et al., 2015).

Revised version: SOs (∼1 Hz) are characterised by fluctuations in neuronal membrane potential between periods of silence (’hyperpolarization’, or ‘DOWN’ state) and excitation (’depolarization’, or ‘UP’ state) (Amzica & Steriade, 2002; Steriade, McCormick, & Sejnowski, 1993). These oscillations originate in neocortical regions, particularly the medial prefrontal cortex (mPFC), during NREM sleep (Dang-Vu et al., 2008; Massimini et al., 2004; Nir et al., 2011). The depolarization phase of SOs is believed to trigger sleep spindles in the thalamus, producing 11–16 Hz oscillations (Fernandez & Lüthi, 2020; Mak-McCully et al., 2017). Thalamic spindles tend to align with the excitable UP-states of SOs and can synchronize with hippocampal ripples (Buzsáki, 2015; Helfrich et al., 2019; Joo & Frank, 2018; Ngo, Fell, & Staresina, 2020). Rodent optogenetics studies have shown that only thalamic spindles phase-locked to the UP-state of cortical SOs enhance memory consolidation, while out-of-phase spindles do not (Latchoumane et al., 2017). This finding is consistent with intracranial evidence from human epilepsy patients, suggesting the thalamus acts as a key relay between the hippocampus and cortex, coordinating their interaction (Coulon, Budde, & Pape, 2012; Ferraris et al., 2021; Schreiner et al., 2022; Staresina et al., 2015).

Earlier version: The triple coupling of SOs, spindles, and ripples is believed to facilitate memory consolidation by synchronizing neuronal activity across brain regions and replaying memory traces from the hippocampus to the neocortex (Diekelmann & Born, 2010; Latchoumane et al., 2017; Liu et al., 2022; Singh, Norman, & Schapiro, 2022; Staresina et al., 2015; Staresina et al., 2023), possibly through relay from thalamic spindles. However, the exact mechanisms of this inter-regional communication during coordinated sleep rhythms remain unclear. Understanding the brain-wide activation associated with these rhythms is crucial, as it reveals how information flows among brain regions, especially during the critical time windows when these rhythms align.

Revised version: The triple coupling of SOs, spindles, and ripples is believed to facilitate memory consolidation by synchronizing neuronal activity across brain regions and replaying memory traces from the hippocampus to the neocortex (Diekelmann & Born, 2010; Latchoumane et al., 2017; Liu et al., 2022; Singh, Norman, & Schapiro, 2022; Staresina et al., 2015; Staresina et al., 2023), possibly through relay from thalamic spindles. Consistent with this view, prior human EEG-fMRI studies have reported spindle-related activation not only in the thalamus, but also in the hippocampus and adjacent parahippocampal gyrus (Bergmann et al., 2012; Schabus et al., 2007). Spindle-related activity has also been linked to striatal engagement, suggesting a broader network that may support memory-related processing during sleep (Fogel et al., 2017). However, the exact mechanisms of this inter-regional communication during coordinated sleep rhythms remain unclear. Understanding the brain-wide activation associated with these rhythms is crucial, as it reveals how information flows among brain regions, especially during the critical time windows when these rhythms align.

Earlier version: Investigating these sleep-related neural processes in humans is challenging because it requires tracking transient sleep rhythms while simultaneously assessing their widespread brain activation. Recent advances in simultaneous EEG-fMRI techniques provide a unique opportunity to explore these processes. EEG allows for precise event-based detection of neural signal, while fMRI provides insight into the broader spatial patterns of brain activation and functional connectivity (Horovitz et al., 2008; Huang et al., 2024; Laufs, 2008; Laufs, Walker, & Lund, 2007; Schabus et al., 2007; Spoormaker et al., 2010). Previous EEG-fMRI studies on sleep have focused on classifying sleep stages or examining the neural correlates of specific waves (Bergmann et al., 2012; Caporro et al., 2012; Czisch et al., 2009; Fogel et al., 2017; Hale et al., 2016; Ilhan-Bayrakcı et al., 2022; Moehlman et al., 2019; Picchioni et al., 2011). These studies have generally reported that slow oscillations are associated with widespread cortical and subcortical BOLD changes, whereas spindles elicit activation in the thalamus, as well as in several cortical and paralimbic regions. Although these findings provide valuable insights into the BOLD correlates of sleep rhythms, they often do not employ sophisticated temporal modeling (Huang et al., 2024), to capture the dynamic interactions between different oscillatory events, e.g., the coupling between SOs and spindles.

Revised version: Investigating these sleep-related neural processes in humans is challenging because it requires tracking transient sleep rhythms while simultaneously assessing their widespread brain activation. Recent advances in simultaneous EEG-fMRI techniques provide a unique opportunity to explore these processes. EEG allows for precise event-based detection of neural signal, while fMRI provides insight into the broader spatial patterns of brain activation and functional connectivity (Horovitz et al., 2008; Huang et al., 2024; Laufs, 2008; Laufs, Walker, & Lund, 2007; Spoormaker et al., 2010). Previous EEG-fMRI studies on sleep have examined both global sleep characteristics (Hale et al., 2016; Moehlman et al., 2019) and the neural correlates of specific waves, including slow oscillations and spindles. These studies have generally shown that slow oscillations are associated with widespread cortical and subcortical BOLD changes (Czisch et al., 2009; Ilhan-Bayrakcı et al., 2022; Picchioni et al., 2011), whereas spindles have been linked not only to thalamic activation but also to cortical and paralimbic regions, including the hippocampus and parahippocampal gyrus (Bergmann et al., 2012; Caporro et al., 2012; Fogel et al., 2017; Schabus et al., 2007). Although these findings provide important insight into the BOLD correlates of sleep rhythms, most previous studies have focused on individual oscillatory events rather than explicitly modelling their temporal interaction (Bergmann et al., 2012; Caporro et al., 2012; Fogel et al., 2017; Picchioni et al., 2011). Only a few recent studies have begun to examine coupling between rhythms directly, for example Huang et al. (2024).

Understanding sleep rhythm-related brain activation is also essential for uncovering its potential cognitive functions. Since performing cognitive tasks during sleep is limited, most neuroimaging studies on human sleep rely on the targeted memory reactivation (TMR) paradigm, where participants are re-exposed to sensory cues during sleep to trigger associated memory reactivation (Cousins et al., 2016; Hu et al., 2020; Oudiette & Paller, 2013; Siefert et al., 2024). For instance, fMRI studies have shown significant hippocampal activation in response to odour re-exposure during sleep (Rasch et al., 2007), while EEG studies suggest that this reactivation is more likely during the UP-state of SOs (Cousins et al., 2016; Göldi et al., 2019; Xia et al., 2023). However, while TMR provides insight into externally induced memory processing during sleep, it offers limited understanding of the endogenous neural mechanisms underlying natural sleep rhythms.

To address this gap, our study investigates brain-wide activation and functional connectivity patterns associated with SO-spindle coupling, and employs a cognitive state decoding approach (Margulies et al., 2016; Yarkoni et al., 2011) - albeit indirectly - to infer potential cognitive functions. In the current study, we used simultaneous EEG-fMRI recordings during nocturnal naps (detailed sleep staging results are provided in the Methods and Table S1) in 107 participants. Although directly detecting hippocampal ripples using scalp EEG or fMRI is challenging, we expected that hippocampal activation in fMRI would coincide with SO-spindle coupling detected by EEG, given that SOs, spindles, and ripples frequently co-occur during NREM sleep. We also anticipated a critical role of the thalamus, particularly thalamic spindles, in coordinating hippocampal-cortical communication.

Earlier version: We found significant coupling between SOs and spindles during NREM sleep (N2/3), with spindle peaks occurring slightly before the SO peak. This coupling was associated with increased activation in both the thalamus and hippocampus, with functional connectivity patterns suggesting thalamic coordination of hippocampal-cortical communication. These findings highlight the key role of the thalamus in coordinating hippocampal-cortical interactions during human sleep and provide new insights into the neural mechanisms underlying sleep-dependent brain communication. A deeper understanding of these mechanisms may contribute to future neuromodulation approaches aimed at enhancing sleep-dependent cognitive function and treating sleep-related disorders.

Revised version: We found significant coupling between SOs and spindles during NREM sleep (N2/3), with spindle peaks occurring slightly before the SO peak. This coupling was associated with increased activation in both the thalamus and hippocampus, with functional connectivity patterns suggesting thalamic coordination of hippocampal-cortical communication, in line with prior EEG-fMRI studies of spindle-related activity (Bergmann et al., 2012; Caporro et al., 2012; Fogel et al., 2017; Schabus et al., 2007). These findings highlight the key role of the thalamus in coordinating hippocampal-cortical interactions during human sleep and provide new insights into the neural mechanisms underlying sleep-dependent brain communication. A deeper understanding of these mechanisms may contribute to future neuromodulation approaches aimed at enhancing sleep-dependent cognitive function and treating sleep-related disorders.

Results

Simultaneous EEG-fMRI data during nocturnal sleep were collected from 138 healthy subjects (age, 22.56 ± 0.30; 81 females). After excluding subjects due to excessive head motion (mean framewise displacement > 0.5) or insufficient scanning sleep duration (<10 min), 107 participants (age, 22.40 ± 0.34; 63 females) were retained for subsequent analysis, with an average scanning time of sleep of 3.01 ± 0.14 hours (see Methods and Table S1). Since NREM sleep predominantly occurs in early sleep (Carskadon & Dement, 2005), we conducted our scans during the first half of the night to ensure sufficient coverage of NREM sleep, this is also the period when sleep rhythms are most consistently observed and studied in relation to memory consolidation (Maingret et al., 2016; Molle et al., 2011; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023).

Sleep stages and sleep rhythms

The primary objective of the EEG analysis was to examine the intrinsic characteristics of sleep rhythms and to identify the timing of these SO and spindle events for subsequent fMRI analyses.

Earlier version: We first removed MRI gradient noise from the EEG signal (detailed in Methods, see Fig. S1) and then applied an automated sleep staging algorithm (Vallat & Walker, 2021). The staging results were manually reviewed by two sleep experts to ensure the accurate classification of each sleep stages (Fig. 1b). Upon validation, we confirmed that N2/3 stages predominated in the dataset, with a significantly higher proportion of N2/3 sleep (75.77% ± 1.50%) compared to N1 (14.57% ± 1.06%, t(106) = 25.04, p < 1e-4, paired-samples t test) and REM sleep (9.66% ± 0.90%, t(106) = 29.60, p < 1e-4, Fig. 1d, Table S1).

Revised version: We first removed MRI gradient noise from the EEG signal (detailed in Methods, see Fig. S1) and then applied an automated sleep staging algorithm (Vallat & Walker, 2021). The staging results were manually reviewed by two sleep experts to ensure the accurate classification of each sleep stage (Fig. 1b). Upon validation, we confirmed that N2/3 stages predominated in the dataset, with a significantly higher proportion of N2/3 sleep (75.77% ± 1.50%) compared to N1 (14.57% ± 1.06%, t(106) = 25.04, p < 1e-4, paired-samples t test) and REM sleep (9.66% ± 0.90%, t(106) = 29.60, p < 1e-4, Fig. 1d, Table S1).

Sleep stages and sleep rhythms in 107 subjects.

Earlier version: a, Sleep rhythms and task schematic. Subjects slept first half of the night with simultaneous EEG-fMRI recordings. Since detecting hippocampal ripples directly from scalp EEG is challenging, our focus was on capturing SOs, spindles, and their couplings. Regions of interest (ROIs) are colour-coded: green for the thalamus (spindle), purple for the mPFC (SOs), and orange for the hippocampus (ripples). b, Sleep staging and EEG spectrogram. N2/3 sleep stages (red line) were initially identified using an offline automatic sleep staging algorithm (Vallat & Walker, 2021) and then manually validated. (c) Schematic of EEG data across different sleep stages, using preprocessed data from the C3 electrode. (d) Proportion of each sleep stage in the dataset. (e) Amplitudes (μV) of SOs (left) and spindles (right) across sleep stages. The SO amplitudes used here are from 0.16-1.25 Hz band EEG data, and the spindle amplitudes used here are from 12-16 Hz band EEG data. Each dot represents an individual participant. Error bars indicate SEM. *** p < 0.001.

Revised version: a, Sleep rhythms and task schematic. Subjects slept first half of the night with simultaneous EEG-fMRI recordings. Since detecting hippocampal ripples directly from scalp EEG is challenging, our focus was on capturing SOs, spindles, and their couplings. Regions of interest (ROIs) are colour-coded: green for the thalamus (spindle), purple for the mPFC (SOs), and orange for the hippocampus (ripples). b, Sleep staging and EEG spectrogram. N2/3 sleep stages (red line) were initially identified using an offline automatic sleep staging algorithm (Vallat & Walker, 2021) and then manually validated. c, Schematic of EEG data across different sleep stages, using preprocessed data from the C3 electrode. d, Proportion of each sleep stage in the dataset. e, Amplitudes (μV) of detected SOs (left) and spindles (right) across sleep stages. SO and spindle detection thresholds were defined from N2/3 sleep within each participant and then applied unchanged to N1 and REM for descriptive comparison. Detections in N1 and REM should therefore be interpreted as detector outputs under this fixed N2/3-derived criterion. The SO amplitudes were measured from the 0.16-1.25 Hz filtered EEG data, and spindle amplitudes were measured from the 12-16 Hz filtered EEG data. Each dot represents an individual participant. Error bars indicate SEM. *** p < 0.001.

Added in the revised version: Each sleep stage is characterised by distinct spectral properties and rhythmic waveforms, serving as physiological markers (Fig. 1c). Because SO and spindle detection relies on amplitude-based percentile thresholds, we avoided estimating separate thresholds within each sleep stage. Instead, for each participant, the SO and spindle thresholds were defined from N2/3 sleep only, where these rhythms are most abundant and physiologically expected, and the same fixed thresholds were then applied to N1 and REM for descriptive comparison.

Added in the revised version: Under this fixed N2/3-derived thresholding, detected SOs and spindles were larger and more frequent in N2/3 than in N1 or REM. SO and spindle amplitudes were significantly higher during N2/3 sleep (SO: 25.59 ± 1.49 μV; spindle: 7.39 ± 0.27 μV) than during N1 (SO: 20.15 ± 2.32 μV; spindle: 5.23 ± 0.27 μV) and REM sleep (SO: 19.84 ± 1.22 μV; spindle: 5.60 ± 0.22 μV; all p < 1e-4; Fig. 1e, Fig. S2). The corresponding event densities showed the same pattern, with 9.64 ± 0.25 SOs/min and 4.19 ± 0.10 spindles/min in N2/3, compared with 2.95 ± 0.16 SOs/min and 2.71 ± 0.14 spindles/min in N1, and 2.07 ± 0.17 SOs/min and 1.81 ± 0.14 spindles/min in REM (all p < 1e-4). We therefore report detections in N1 and REM only as descriptive outputs of the detector under a fixed N2/3-derived criterion, rather than as physiological equivalents of canonical N2/3 SOs or spindles.

Earlier version: Each sleep stage is characterized by distinct spectral properties and rhythmic waveforms, serving as physiological markers (Fig. 1c). As a sanity check, we compared SO and spindle amplitudes and densities across stages. Both were significantly higher during N2/3 sleep (SO: 25.59 ± 1.49 μV, spindle: 7.39 ± 0.27 μV) compared to N1 (SO: 18.87 ± 1.08 μV, spindle: 6.72 ± 0.25 μV) and REM (SO: 16.67 ± 1.14 μV, spindle: 6.23 ± 0.23 μV, all p < 1e-4, Fig. 1e, Fig. S2). This aligns with the established understanding that N2/3 stages feature more pronounced slow-wave activity and robust spindles (Gaillard & Blois, 1981). Furthermore, we found that the amplitude of the SO DOWN-state (13.62 ± 0.67 μV) was significantly higher than the UP-state (12.98 ± 0.66 μV, t(106) = 4.82, p < 1e-4), consistent with previous work (Cash et al., 2009; Dang-Vu et al., 2008). Based on this, we used the DOWN-state to lock the timing of SOs for subsequent EEG-informed fMRI analysis. Descriptive information on sleep rhythms during these stages for all 107 subjects can be found in Table S2-S4.

Revised version: This aligns with the established understanding that N2/3 stages feature more pronounced slow-wave activity and robust spindles (Gaillard & Blois, 1981). Furthermore, we found that the amplitude of the SO DOWN-state (13.62 ± 0.67 μV) was significantly higher than the UP-state (12.98 ± 0.66 μV, t(106) = 4.82, p < 1e-4), consistent with previous work (Cash et al., 2009; Dang-Vu et al., 2008). Based on this, we used the DOWN-state to lock the timing of SOs for subsequent EEG-informed fMRI analysis. Descriptive information on sleep rhythms during these stages for all 107 subjects can be found in Table S2-S4.

SO-spindle coupling

Added in the revised version: SO-spindle coupling is considered important for sleep-dependent memory consolidation. In the current study, using the same N2/3-derived detection thresholds described above, we found that SO-spindle coupling occurred most frequently during N2/3 sleep (2.46 ± 0.06 events/min). Coupling density was significantly lower in N1 (0.75 ± 0.05 events/min, t(106) = 23.54, p < 1e-4) and REM sleep (0.43 ± 0.04 events/min, t(106) = 31.24, p < 1e-4; Fig. 2b, Table S2-S4), consistent with the expected predominance of SO-spindle coupling in NREM sleep (Ngo et al., 2013; Staresina et al., 2015). As with the individual SO and spindle detections, coupling events detected in N1 and REM were retained only for descriptive stage-wise reporting (see Table S2, S4). They were not used to support physiological claims about SO-spindle coupling in these stages, and they were not entered into the EEG-informed fMRI analyses. All subsequent fMRI GLM and PPI analyses were restricted to N2/3 sleep.

Removed in the revised version: SO-spindle coupling is considered important for sleep-dependent memory consolidation. In current study, we found that this coupling primarily occurs during N2/3 sleep stages (2.46 ± 0.06 times/min), significantly more than in N1 (2.23 ± 0.09 times/min, t(106) = 2.81, p = 0.0059) and REM stages (2.32 ± 0.09 times/min, t(106) = 1.99, p = 0.0492, Fig. 2b, Table S2-S4), consistent with previous findings (Ngo et al., 2013; Staresina et al., 2015). After extracting all EEG epochs where SO-spindle coupling occurred, we analysed their spectral and phase characteristics. The spindles were most likely to occur slightly before the UP state peak of SOs (Fig. 2a, e), aligning with results from both animal studies (Maingret et al., 2016) and human research (Staresina et al., 2015). In our data, this pattern was consistent across subjects (Fig. 2d, Rayleigh test: z = 9.51, p < 1e-4), with the peak of the spindle aligned at an SO phase of −41.61 ± 0.86° (the SO UP state peak is 0°).

Sleep rhythms and SO-spindle coupling.

Earlier version: a, The SO-spindle coupling in the temporal frequency domain. The upper two rows illustrates the spindle (12-16 Hz) phase-locked in the transition to UP-state of SO (0.16-1.25 Hz). The bottom row shows the averaged temporal frequency pattern across all instances of SO-spindle coupling and over all subjects. b, SO-spindle coupling density across different sleep stages. c, Differences between coupled and uncoupled sleep rhythms. The left panel shows the difference in amplitude between spindles coupled with SOs (Coupling) and spindles not coupled with SOs (Other). The right panel displays the difference in amplitude between SOs coupled with spindles (Coupling) and SOs not coupled with spindles (Other). d, Phase modulation of SO-spindle coupling. Spindle peaks cluster slightly before the UP-state peak of SO (i.e., 0°), where –π/2 reflects the transition from DOWN to UP-state. The histogram represents the distribution of coupling directions across all subjects, with the red line showing the mean. Coupling phases for each subject are plotted on the circle, with coupling strength color-coded. e, Distribution of spindle peaks on the SO phase during all SO-spindle coupling events across participants. The distribution is represented by a probability density function, and the density is evaluated at 100 equally spaced points covering the data range. Each dot represents data from an individual subject. Error bars indicate the SEM. * p < 0.05, ** p < 0.01, *** p < 0.001.

Revised version: a, The SO-spindle coupling in the temporal frequency domain. The upper two rows illustrates the spindle (12-16 Hz) phase-locked in the transition to UP-state of SO (0.16-1.25 Hz). The bottom row shows the averaged temporal frequency pattern across all instances of SO-spindle coupling and over all subjects. b, SO-spindle coupling density across sleep stages, using SO and spindle detections obtained with fixed N2/3-derived thresholds. Coupling events in N1 and REM are shown only for descriptive comparison. The EEG-informed fMRI analyses used N2/3 coupling events only. c, Differences between coupled and uncoupled sleep rhythms. The left panel shows the difference in amplitude between spindles coupled with SOs (Coupling) and spindles not coupled with SOs (Other). The right panel displays the difference in amplitude between SOs coupled with spindles (Coupling) and SOs not coupled with spindles (Other). d, Phase modulation of SO-spindle coupling. Spindle peaks cluster slightly before the UP-state peak of SO (i.e., 0°), where –π/2 reflects the transition from DOWN to UP-state. The histogram represents the distribution of coupling directions across all subjects, with the red line showing the mean. Coupling phases for each subject are plotted on the circle, with coupling strength colour-coded. e, Distribution of spindle peaks on the SO phase during all SO-spindle coupling events across participants. The distribution is represented by a probability density function, and the density is evaluated at 100 equally spaced points covering the data range. Each dot represents data from an individual subject. Error bars indicate the SEM. *** p < 0.001.

Added in the revised version: After extracting all N2/3 EEG epochs in which SO-spindle coupling occurred, we analysed their spectral and phase characteristics. The spindles were most likely to occur slightly before the UP-state peak of SOs (Fig. 2a, e), aligning with results from both animal studies (Maingret et al., 2016) and human research (Staresina et al., 2015). In our data, this pattern was consistent across subjects (Fig. 2d, Rayleigh test: z = 9.51, p < 1e-4), with the peak of the spindle aligned at an SO phase of −41.61 ± 0.86° (the SO UP-state peak is 0°).

Interestingly, spindle events coupled with SOs had higher amplitudes (7.69 ± 0.26 μV) compared to those not coupled (7.33 ± 0.25 μV, t(106) = 4.65, p < 1e-4, Fig. 2c, left). Similarly, SO events coupled with spindles exhibited higher amplitudes (26.28 ± 1.34 μV) than those not coupled (24.56 ± 1.33 μV, t(106) = 6.86, p < 1e-4, Fig. 2c, right). These findings suggest coordinated neural processing during SO-spindle coupling. Next, we explore the brain-wide activation associated with these sleep rhythms.

Sleep rhythms associated brain activation

Earlier version: To investigate brain-wide activity during sleep rhythms, we modeled the fMRI blood oxygen level-dependent (BOLD) signal based on EEG-identified timing of SOs and spindle events during N2/3 sleep stages. The design matrix included the main effects of SOs, spindles, and their coupling (Fig. 3a, detailed in Methods). This method is validated in our previous work to effectively capture replay-aligned brain activation during rest (Huang et al., 2024). We found that SOs were associated with positive activation in the thalamus (ROI analysis, t(106) = 2.38, p = 0.0096, one-sample t test) and significant deactivation in the neocortex (Fig. 3b, see also Fig. S5), particularly in the default mode network (DMN), which includes the mPFC and posterior cingulate cortex (PCC), consistent with previous research (Dang-Vu et al., 2008; Picchioni et al., 2011; Tagliazucchi et al., 2013). This deactivation of the neocortex, particularly the DMN, at the trough of SOs is intriguing, given that the DMN is thought to encode the internal model of the world or cognitive map (Baldassano, Hasson, & Norman, 2018; Constantinescu, O’Reilly, & Behrens, 2016; Park, Miller, & Boorman, 2021).

Revised version: To investigate brain-wide activity during sleep rhythms, we modelled the fMRI blood oxygen level-dependent (BOLD) signal based on EEG-identified timing of SOs and spindle events during N2/3 sleep stages. The design matrix included the main effects of SOs, spindles, and their coupling (Fig. 3a, detailed in Methods). This method is validated in our previous work to effectively capture replay-aligned brain activation during rest (Huang et al., 2024). We found that SOs were associated with positive activation in the thalamus (ROI analysis, t(106) = 2.38, p = 0.0096, one-sample t test) and significant deactivation in the neocortex (Fig. 3b, see also Fig. S5), particularly in the default mode network (DMN), which includes the mPFC and posterior cingulate cortex (PCC), consistent with previous research (Dang-Vu et al., 2008; Picchioni et al., 2011; Tagliazucchi et al., 2013). This deactivation of the neocortex, particularly the DMN, at the trough of SOs is intriguing, given that the DMN is thought to encode the internal model of the world or cognitive map (Baldassano, Hasson, & Norman, 2018; Constantinescu, O’Reilly, & Behrens, 2016; Park, Miller, & Boorman, 2021).

Brain-wide activation associated with sleep rhythms.

a, Simultaneous EEG-fMRI analysis framework for detecting brain-wide activation during sleep rhythms. Detected SOs, spindles, and their coupling were convolved with the hemodynamic response function (HRF) and downsampled to match fMRI temporal resolution. These events formed the design matrix for the general linear model (GLM) analysis of fMRI activity during sleep, linking the EEG-derived timing of sleep rhythms to the corresponding brain responses in fMRI. b, Brain-wide activation associated with SOs. The upper row illustrates SOs, and the lower row shows the fMRI activation pattern during SO events, whole-brain family-wise error (FWE) corrected at the cluster level (p < 0.05) with a cluster-forming voxel threshold of p < 0.001. c, Brain-wide activation associated with spindles. Same as panel b, but for spindle events. d, Brain-wide activation associated with SO-spindle coupling (compared to non-coupling events).

Earlier version: Spindles were correlated with positive activation in the thalamus (ROI analysis, t(106) = 15.39, p < 1e-4), the anterior cingulate cortex (ACC), and the putamen, alongside deactivation in the DMN (Fig. 3c). Notably, SO-spindle coupling was linked to significant activation in both the thalamus (ROI analysis, t(106) = 3.38, p = 0.0005) and the hippocampus (ROI analysis, t(106) = 2.50, p = 0.0070, Fig. 3d). However, no decrease in DMN activity was found during SO-spindle coupling, and DMN activity during SO was significantly lower than during coupling (ROI analysis, t(106) = -4.17, p < 1e-4). For more detailed activation patterns, see Table S5-S7. We also varied the threshold used to detect SO events to assess its effect on hippocampal activation during SO-spindle coupling and observed that hippocampal activation remained significant when the percentile thresholds for SO detection ranged between 71% and 80% (see Fig. S6).

Revised version: Spindles were correlated with positive activation in the thalamus (ROI analysis, t(106) = 15.39, p < 1e-4), the anterior cingulate cortex (ACC), and the putamen, alongside deactivation in the DMN (Fig. 3c). Notably, SO-spindle coupling was linked to significant activation in both the thalamus (ROI analysis, t(106) = 3.38, p = 0.0005) and the hippocampus (ROI analysis, t(106) = 2.50, p = 0.0070, Fig. 3d). However, no decrease in DMN activity was found during SO-spindle coupling, and DMN activity during SO was significantly lower than during coupling (ROI analysis, t(106) = −4.17, p < 1e-4). For more detailed activation patterns, see Table S5-S7.

Added in the revised version: To ensure the results were not driven by individual differences or parameter selection, we conducted a series of control analyses. First, we excluded participants with less than 10 minutes of N2/3 sleep. Only three participants met this criterion, and their exclusion did not change the main results. For example, hippocampal activation during SO-spindle coupling remained significant (t(103) = 2.50, p = 0.0071), comparable to the full sample (t(106) = 2.50, p = 0.0070). Second, because the absolute number of detected SO-spindle coupling events depends on the SO detection threshold, we examined whether the main EEG-fMRI results were sensitive to this parameter. To this end, we varied the SO percentile threshold and reconstructed the EEG-informed GLM at each level. Hippocampal activation during SO-spindle coupling remained significant across a range of thresholds (71st - 80th percentile; Fig. S6). Third, to test whether the results depended on the use of a single lateralised frontal electrode, we repeated the EEG-informed fMRI GLM using events detected from Fz. Hippocampal activation during SO-spindle coupling again remained significant (t(106) = 2.47, p = 0.0076), closely matching the original F3-based result (t(106) = 2.50, p = 0.0070).

To understand the functional implications of these sleep rhythm-related brain activation, we performed open-ended cognitive state decoding using the NeuroSynth database (Yarkoni et al., 2011), examining associations between topic terms and regions of interest derived from the activation patterns of the sleep rhythms. Topic terms were ranked by their z-scores across the rhythms to assess overall functional relevance (see Methods for details). For positive activation patterns (Fig. S7a), SOs were associated with declarative (z = 6.31, p < 1e-4) and episodic memory (z = 5.08, p < 1e-4), while spindles aligned most with working memory (z = 4.08, p < 1e-4). SO-spindle coupling showed the strongest similarity to episodic (z = 3.66, p = 0.0003) and declarative memory (z = 2.21, p = 0.0271). For negative (deactivation) patterns (Fig. S7b), SOs were linked to declarative memory (z = 7.61, p < 1e-4), while spindles resembled task-related activations (z = 7.05, p < 1e-4). SO-spindle coupling had the highest similarity to task-related (z = 9.40, p < 1e-4) and working memory (z = 9.15, p < 1e-4). These findings suggest that internal memory-related regions are more likely to reactivate during SO-spindle coupling, whereas task-related activations, like working memory, are less prominent.

Functional connectivity during SO-spindle coupling

Previous studies have indicated that the hippocampal-cortical dialogue during sleep memory consolidation relies on the thalamus (Coulon, Budde, & Pape, 2012; Ferraris et al., 2021; Latchoumane et al., 2017). To explore this in the human brain, we examined how functional connectivity (FC) changes when spindles couple (versus not) with SOs during sleep, using psychophysiological interaction (PPI) analysis (Friston et al., 1997). Building on the previous GLM (Fig. 3a), this model incorporated two additional regressors: the BOLD signal of the ROI seed and its interaction with SO-spindle events (Fig. 4a, detailed in Methods).

Functional connectivity changes during SO-spindle coupling.

a, The PPI analysis framework for detecting brain-wide connectivity changes during SO-spindle coupling. This starts by setting a specific ROI (e.g., the hippocampus) as the seed to extract the BOLD signal (physiological condition) and using identified SO-spindle coupling events as the psychological condition to compute the interaction term. The design matrix includes the main effects of the physiological and psychological conditions, along with their interaction. This analysis examines whether whole-brain communication with the hippocampus changes as a function of SO-spindle coupling. b, Hippocampus-based functional connectivity with the whole brain (main effect of hippocampus BOLD signal in PPI analysis). The hippocampus ROI is bilateral, anatomically defined (bottom, orange colour). Brain-wide connectivity is shown with whole-brain FWE correction at the cluster level (p < 1e-7) with a cluster-forming voxel threshold of punc. < 0.001 for visualization purpose. c, Same with panel b, but based on thalamus (bilateral anatomically defined ROI). d, Same with panel b, but based on the mPFC (bilateral functionally defined ROI, detailed in Methods). e, Functional connectivity changes during SO-spindle coupling for hippocampus-based (left bottom, orange colour), thalamus-based (middle, green colour), and mPFC-based (right bottom, purple colour) connectivity. The results of ROI analysis for each direction are shown on the arrows. * p < 0.05, ns., not significant. Abbreviations: FC - functional connectivity, PPI - psychophysiological interaction.

We first examined whole-brain FC (physiological term) in three ROIs: the hippocampus (bilateral, anatomically defined), thalamus (bilateral, anatomically defined), and mPFC (bilateral, functionally defined). The hippocampus-based FC was mainly connected with the DMN, particularly the mPFC (Fig. 4b). The thalamus-based FC was primarily linked to the basal ganglia and ACC (Fig. 4c), and the mPFC-based FC showed strong connections with the hippocampus and PCC (Fig. 4d).

Crucially, the interaction between FC and SO-spindle coupling revealed that only the functional connectivity of hippocampus -> thalamus (ROI analysis, t(106) = 1.86, p = 0.0328) and thalamus -> mPFC (ROI analysis, t(106) = 1.98, p = 0.0251) significantly increased during SO-spindle coupling, with no significant changes in all other pathways (Fig. 4e). We also conducted PPI analyses for the other two events (SOs and spindles), and neither yielded significant connectivity changes in the three ROIs, as all failed to survive whole-brain FWE correction at the cluster level (p < 0.05). Together, these findings suggest that the thalamus, likely via spindles, coordinates hippocampal-cortical communication selectively during SO-spindle coupling, but not isolated SOs or spindle events alone.

Discussion

Using simultaneous EEG and fMRI during nocturnal naps in 107 subjects, we identified robust coupling between spindles and the UP-state of SOs in deep NREM sleep (N2/3 stages). This coupling was linked to distinct neural activation patterns, including significant thalamic activation during spindle occurrences, reduced activity in the DMN during SOs trough, and selective hippocampal activation when spindles and SOs coupled. Furthermore, functional connectivity analyses revealed increased connectivity from the hippocampus to the thalamus and from the thalamus to the mPFC, only during SO-spindle coupling, but not isolated SOs or spindle events alone. These results shed light on how the thalamus may coordinate hippocampal-cortical interactions during NREM sleep.

Our data provide insights into the neurobiological underpinnings of these sleep rhythms. SOs, originating mainly in neocortical areas such as the mPFC, alternate between DOWN- and UP-states. The thalamus generates sleep spindles, which in turn couple with SOs. Our finding that spindle peaks consistently occurred slightly before the UP-state peak of SOs (in 83 out of 107 participants), concurs with prior studies, including Schreiner et al. (2021). Yet it differs from some results suggesting spindles might peak right at the SO UP-state (Hahn et al., 2020; Helfrich et al., 2018). Such discrepancies could arise from differences in detection algorithms, participant age (Helfrich et al., 2018), or subtle variations in cortical-thalamic timing. Nonetheless, these results underscore the importance of coordinated SO-spindle interplay in supporting sleep-dependent processes.

Earlier version: When modeling the timing of these sleep rhythms in the fMRI, we observed hippocampal activation selectively during SO-spindle events. This suggests the possibility of triple coupling (SOs–spindles– ripples), even though our scalp EEG was not sufficiently sensitive to detect hippocampal ripples—key markers of memory replay (Buzsáki, 2015). Recent iEEG evidence indicates that ripples often co-occur with both spindles (Ngo, Fell, & Staresina, 2020) and SOs (Staresina et al., 2015; Staresina et al., 2023). Therefore, the hippocampal involvement during SO-spindle events in our study may reflect memory replay from the hippocampus, propagated via thalamic spindles to distributed cortical regions.

Revised version: When modelling the timing of these sleep rhythms in the fMRI, we observed hippocampal activation selectively during SO-spindle events. This suggests the possibility of triple coupling (SOs–spindles–ripples), even though our scalp EEG was not sufficiently sensitive to detect hippocampal ripples—key markers of memory replay (Buzsáki, 2015). Recent iEEG evidence indicates that ripples often co-occur with both spindles (Ngo, Fell, & Staresina, 2020) and SOs (Staresina et al., 2015; Staresina et al., 2023). Therefore, the hippocampal involvement during SO-spindle events in our study may reflect memory replay from the hippocampus, propagated via thalamic spindles to distributed cortical regions.

The thalamus, known to generate spindles (Halassa et al., 2011), plays a key role in producing and coordinating sleep rhythms (Coulon, Budde, & Pape, 2012; Crunelli et al., 2018), while the hippocampus is found essential for memory consolidation (Buzsáki, 2015; Diba & Buzsáki, 2007; Singh, Norman, & Schapiro, 2022). The increased hippocampal and thalamic activity, along with strengthened connectivity between these regions and the mPFC during SO-spindle events, underscores a hippocampal-thalamic-neocortical information flow. This aligns with recent findings suggesting the thalamus orchestrates neocortical oscillations during sleep (Schreiner et al., 2022). The thalamus and hippocampus thus appear central to memory consolidation during sleep, guiding information transfer to the neocortex, e.g., mPFC.

Earlier version: An intriguing aspect of our findings is the reduced DMN activity during SOs when modeled at the SO trough (DOWN-state). This reduced DMN activity may reflect large-scale neural inhibition characteristic of the SO trough. The DMN is typically active during internally oriented cognition (e.g., self-referential processing or mind-wandering) and is suppressed during external stimuli processing (Yeshurun, Nguyen, & Hasson, 2021). It is unlikely, however, that this suppression of DMN during SO events is related to a shift from internal cognition to external responses given it is during deep sleep time. Instead, it could be driven by the inherent rhythmic pattern of SOs, which makes it difficult to separate UP-from DOWN-states (the two temporal regressors were highly correlated, and similar brain activation during SOs events was obtained if modelled at the SO peak instead, Fig. S5). Since the amplitude at the SO trough is consistently larger than that at the SO peak, the neural activation we detected may primarily capture the large-scale inhibition from DOWN-state. Interestingly, no such DMN reduction was found during SO-spindle coupling, implying that coupling may involve distinct neural dynamics that partially re-engage DMN-related processes, possibly reflecting memory-related reactivation. Future research using high-temporal-resolution techniques like iEEG could clarify these possibilities.

Revised version: An intriguing aspect of our findings is the reduced DMN activity during SOs when modelled at the SO trough (DOWN-state). This reduced DMN activity may reflect large-scale neural inhibition characteristic of the SO trough. The DMN is typically active during internally oriented cognition (e.g., self-referential processing or mind-wandering) and is suppressed during external stimuli processing (Yeshurun, Nguyen, & Hasson, 2021). It is unlikely, however, that this suppression of DMN during SO events is related to a shift from internal cognition to external responses given it is during deep sleep time. Instead, it could be driven by the inherent rhythmic pattern of SOs, which makes it difficult to separate UP- from DOWN-states (the two temporal regressors were highly correlated, and similar brain activation during SOs events was obtained if modelled at the SO peak instead, Fig. S5). Since the amplitude at the SO trough is consistently larger than that at the SO peak, the neural activation we detected may primarily capture the large-scale inhibition from DOWN-state. Interestingly, no such DMN reduction was found during SO-spindle coupling, implying that coupling may involve distinct neural dynamics that partially re-engage DMN-related processes, possibly reflecting memory-related reactivation. Future research using high-temporal-resolution techniques like iEEG could clarify these possibilities.

To explore functional relevance, we employed an open-ended cognitive state decoding approach using meta-analytic data (NeuroSynth: Yarkoni et al. (2011)). Although this method usefully generates hypotheses about potential cognitive processes, particularly in the absence of a pre- and post-sleep memory task, it is inherently indirect. Many cognitive terms showed significant associations (16 of 50), such as “episodic memory,” “declarative memory,” and “working memory.” We focused on episodic/declarative memory given the known link with hippocampal reactivation (Diekelmann & Born, 2010; Staresina et al., 2015; Staresina et al., 2023). Nonetheless, these inferences regarding memory reactivation should be interpreted cautiously without direct behavioral measures. Future research incorporating explicit tasks before and after sleep would more rigorously validate these potential functional claims.

Earlier version: Despite providing new insights, our study has several limitations. First, our scalp EEG did not directly capture hippocampal ripples, preventing us from conclusively demonstrating triple coupling. Second, the combination of EEG-fMRI and the lack of a memory task limit our ability to parse fine-grained BOLD responses at the DOWN- vs. UP-states of SOs and link observed activations to behavioral outcomes. Third, the use of large anatomical ROIs may mask subregional contributions of specific thalamic nuclei or hippocampal subfields. Finally, without a memory task, we cannot establish a direct behavioral link between sleep-rhythm-locked activation and memory consolidation. Future studies combining techniques such as ultra-high-field fMRI or iEEG with cognitive tasks may refine our understanding of subregional network dynamics and functional significance during sleep.

Revised version: Despite providing new insights, our study has several limitations. First, our scalp EEG did not directly capture hippocampal ripples, preventing us from conclusively demonstrating triple coupling. Second, the combination of EEG-fMRI and the lack of a memory task limit our ability to parse fine-grained BOLD responses at the DOWN- vs. UP-states of SOs and link observed activations to behavioral outcomes. Third, sleep oscillation detection was based on a single frontal electrode. This choice improved signal stability and event timing in the prolonged simultaneous EEG-fMRI setting, but it did not exploit the full multichannel EEG information and cannot characterise the full spatial distribution of SOs and spindles. In particular, F3-based detection may be more sensitive to frontal SOs and frontal sigma activity than to the centro-parietal fast spindle component. We therefore interpret the EEG-informed fMRI results as reflecting BOLD activity associated with frontal-channel detected SOs, spindles, and their coupling during N2/3 sleep. Future studies using multichannel or source-informed detection strategies, with separate treatment of slow and fast spindles, will be better suited to capture the spatial dynamics of these sleep oscillations. Fourth, the use of large anatomical ROIs may mask subregional contributions of specific thalamic nuclei or hippocampal subfields. Finally, without a memory task, we cannot establish a direct behavioral link between sleep-rhythm-locked activation and memory consolidation. Future studies combining ultra-high-field fMRI or iEEG with cognitive tasks, as well as multichannel or source-informed detection strategies that separately characterize slow and fast spindles, will be better suited to refine our understanding of subregional network dynamics and the functional significance of sleep oscillations.

In summary, we found that SO-spindle coupling during deep NREM sleep is associated with distinct activation in both the thalamus and hippocampus, accompanied by increased connectivity from the hippocampus to the thalamus and from the thalamus to the mPFC. These results emphasize the central role of the thalamus in orchestrating hippocampal-cortical interactions and highlight potential neural substrates underlying the complex interplay of SOs, spindles, and, possibly, ripples. Our open-ended decoding analyses suggest that these coordinated sleep events may be relevant for memory-related processes; however, future experiments incorporating explicit behavioral measures are necessary to confirm their functional significance. By revealing how these rhythms modulate widespread brain networks, our findings open avenues for developing targeted neuromodulatory strategies to enhance sleep-dependent cognitive functions and potentially address sleep-related disorders.

Methods

Participants & Protocol

A total of 138 healthy adults (age: 22.56 ± 0.30; 81 females) were recruited for the study. All participants had either normal vision or vision corrected to normal standards. None had any past psychiatric or neurological disorders. Prior to participation, they were screened for eligibility to undergo magnetic resonance imaging (MRI). The experiment received approval from Peking University ethics committee (reference number: #2015-26, #2018-11-05). Every participant provided written informed consent.

At the start of the experiment, subjects arrived early at the laboratory for registration. They signed informed consent forms, washed their hair, changed their clothes, and put on the EEG cap in preparation for the experiment. After these preliminary steps, participants entered the MRI scanning room by around 20:30, settled down for sleep, and data collection began. Given concern for potential data drift and decline in the signal-to-noise ratio due to extended and continuous fMRI scanning, the data acquisition was segmented into several sessions for data collection. A “sleep” session was concluded either when a participant was fully awake (unable to fall asleep again or needed to use the restroom) or after 8192 seconds (the maximum scanning time of the scan, equating to 4096 volumes with 2 seconds TR).

After removing subjects with excessive head movements (mean FD > 0.5) or those whose sleep durations were under 10 minutes, the final participant count for further analyses was 107 (age: 22.40 ± 0.34; 63 females) with a mean scanning duration of sleep of 3.01 ± 0.14 hours.

EEG data acquisition

EEG was recorded simultaneously with fMRI data using an MR-compatible EEG amplifier system (BrainAmps MR-Plus, Brain Products, Germany), along with a specialized electrode cap. The recording was done using 64 channels in the international 10/20 system, with the reference channel positioned at FCz. In order to adhere to polysomnography (PSG) recording standards, six electrodes were removed from the EEG cap: one for electrocardiogram (ECG) recording, two for electrooculogram (EOG) recording, and three for electromyogram (EMG) recording. EEG data was recorded at a sample rate of 5000 Hz, the resistance of the reference and ground channels was kept below 10 kΩ, and the resistance of the other channels was kept below 20 kΩ. To synchronize the EEG and fMRI recordings, the BrainVision recording software (BrainProducts, Germany) was utilized to capture triggers from the MRI scanner. The EEG sampling was synchronized to the MR scanner’s 10 MHz gradient system clock, ensuring a stable gradient artifact shape over time and enabling accurate artifact removal. This was achieved via the standard clock synchronization interface of the EEG amplifier, minimizing timing jitter and drift.

MRI data acquisition

All MRI data were acquired using a 20-channel head coil on a research-dedicated 3-Tesla Siemens Magnetom Prisma MRI scanner. Earplugs and cushions were provided for noise protection and head motion restriction. We chose the 20-channel head coil because it was compatible with our EEG system and ensured sufficient signal-to-noise ratio (SNR) for both EEG and fMRI acquisition. The EEG electrode cables were guided through the lateral and posterior openings of the head coil, secured with foam padding to reduce motion and minimize MR-related artifacts. Moreover, given the extended nature of nocturnal sleep recordings, the 20-channel coil helped maintain participant comfort while still achieving high-quality simultaneous EEG-fMRI data.

For registration purposes, high-resolution T1-weighted anatomical images were acquired using a three-dimensional magnetization prepared rapid acquisition gradient-echo sequence (sequence specification: number of slices = 192, repetition time (TR) = 2,530 ms, echo time (TE) = 2.98 ms, inversion time = 1,100 ms, voxel size = 0.5 × 0.5 × 1 mm3, flip angle (FA) = 7°). The participants were asked to lie quietly in the scanner during data acquisition.

Earlier version: Then, the “sleep” session began after the participants were instructed to try and fall asleep. For the functional scans, whole-brain images were acquired using k-space and steady-state T2*-weighted gradient echo-planar imaging (EPI) sequence that is sensitive to the BOLD contrast. This measures local magnetic changes caused by changes in blood oxygenation that accompany neural activity (sequence specification: 33 slices in interleaved ascending order, TR = 2000 ms, TE = 30 ms, voxel size = 3.5 × 3.5 × 4.2 mm3, FA = 90°, matrix = 64 × 64, gap = 0.7 mm). A relatively long TR and larger voxel size were chosen to optimize SNR and reduce susceptibility-related distortions, which are critical in EEG-fMRI sleep studies where head motion and physiological noise can be substantial. The longer TR allowed whole-brain coverage with sufficient temporal resolution, while the larger voxel size helped enhance BOLD sensitivity and minimize partial volume effects in deep brain structures (e.g., the thalamus and hippocampus), which are key regions of interest in this study.

Revised version: Then, the “sleep” session began after the participants were instructed to try and fall asleep. For the functional scans, whole-brain images were acquired using a T2*-weighted gradient echo-planar imaging (EPI) sequence sensitive to the BOLD contrast. The sequence parameters were as follows: 33 slices in interleaved ascending order, TR = 2000 ms, TE = 30 ms, voxel size = 3.5 × 3.5 × 4.2 mm3, FA = 90°, matrix = 64 × 64, gap = 0.7 mm. A relatively large voxel size was chosen to preserve signal-to-noise ratio while maintaining whole-brain coverage within a feasible repetition time. This compromise was important for the prolonged overnight EEG-fMRI sleep protocol, where head motion, physiological noise, participant comfort, and sustained acquisition stability are substantial practical constraints (Bodurka et al., 2007; Laufs et al., 2008). A smaller voxel size would have improved spatial specificity, but would have required either a longer repetition time, reduced brain coverage, or lower signal-to-noise ratio.

EEG data preprocessing and sleep stage scoring

EEG data collected inside MRI scanner was contaminated by imaging, ballistocardiographic (BCG) and ocular artifacts. To deal with these artifacts, the preprocessing processes was divided into two parts: automated processing and manual processing.

The automated processing was carried out using custom Python scripts and the MNE toolkit (Gramfort et al., 2013). This encompassed the following procedures: I. The raw EEG data were downsampled from 5000 Hz to 2048 Hz. This downsampled rate accelerates subsequent analyses while retaining the necessary precision to distinguish MR noise. II. Imaging artifacts were removed using the Average Artefact Subtraction method (Allen, Josephs, & Turner, 2000). III. We then applied 0.1 - 30 Hz band-pass finite impulse response (FIR) filters to all channels. The EEG and EOG electrodes were re-referenced by subtracting the value from the M1 or M2 electrode (whichever was from the opposite hemisphere). The outer EMG electrodes were re-referenced to the middle one. IV. The signal space projection algorithm was applied to extract all epochs based on TRs and remove residual MR noise, as well as automatically extracting all heartbeat epochs using ECG electrode data for the removal of BCG noise. V. The data were further downsampled to 100 Hz and subjected to Independent Component Analysis (ICA).

Following this, the manual processing is carried out with labelling and removing the physiological artifacts ICs belonging to eye, muscle and residual BCG through the use of custom MATLAB scripts and the EEGLAB toolkit (Delorme & Makeig, 2004). Sleep stages were then scored by YASA toolkit (Vallat & Walker, 2021) after preprocessing, necessitating data from electrodes C3, left EOG, and left EMG. The classification of N2/3 stages were further verified by experts for validity.

MRI data preprocessing

Results included in this manuscript come from preprocessing performed using fMRIPrep 21.0.4 (Esteban et al., 2020; Esteban et al., 2019), which is based on Nipype 1.6.1 (Gorgolewski et al., 2011). Many internal operations of fMRIPrep use Nilearn 0.8.1 (Abraham et al., 2014), mostly within the functional processing workflow. For more details of the pipeline, see https://fmriprep.readthedocs.io/en/latest/workflows.html.

Removed in the revised version: Conversion of data to the brain imaging data structure standard

Earlier version: To facilitate further analysis and sharing of data, all study data were arranged according to the Brain Imaging Data Structure (BIDS) specification using dcm2bids tool, which is freely available from https://unfmontreal.github.io/Dcm2Bids/.

Revised version: Conversion of data to the brain imaging data structure standard. To facilitate further analysis and sharing of data, all study data were arranged according to the Brain Imaging Data Structure (BIDS) specification using dcm2bids tool, which is freely available from https://unfmontreal.github.io/Dcm2Bids/.

Removed in the revised version: Anatomical data preprocessing

Earlier version: The T1-weighted (T1w) image was corrected for intensity non-uniformity (INU) with N4BiasFieldCorrection (Tustison et al., 2010), distributed with ANTs 2.3.3 (Avants et al., 2008), and used as T1w-reference throughout the workflow. The T1w-reference was then skull-stripped with a Nipype implementation of the antsBrainExtraction.sh workflow (from ANTs), using OASIS30ANTs as target template. Brain tissue segmentation of cerebrospinal fluid (CSF), white-matter (WM) and gray-matter (GM) was performed on the brain-extracted T1w using fast (FSL 6.0.5.1:57b01774, RRID:SCR_002823) (Zhang, Brady, & Smith, 2001). Volume-based spatial normalization to one standard space (MNI152NLin2009cAsym) was performed through nonlinear registration with antsRegistration (ANTs 2.3.3), using brain-extracted versions of both T1w reference and the T1w template. The following template was selected for spatial normalization: ICBM 152 Nonlinear Asymmetrical template version 2009c (Fonov, 2009) [RRID:SCR_008796; TemplateFlow ID: MNI152NLin2009cAsym].

Revised version: Anatomical data preprocessing. The T1-weighted (T1w) image was corrected for intensity non-uniformity (INU) with N4BiasFieldCorrection (Tustison et al., 2010), distributed with ANTs 2.3.3 (Avants et al., 2008), and used as T1w-reference throughout the workflow. The T1w-reference was then skull-stripped with a Nipype implementation of the antsBrainExtraction.sh workflow (from ANTs), using OASIS30ANTs as target template. Brain tissue segmentation of cerebrospinal fluid (CSF), white-matter (WM) and gray-matter (GM) was performed on the brain-extracted T1w using fast (FSL 6.0.5.1:57b01774, RRID:SCR_002823) (Zhang, Brady, & Smith, 2001). Volume-based spatial normalization to one standard space (MNI152NLin2009cAsym) was performed through nonlinear registration with antsRegistration (ANTs 2.3.3), using brain-extracted versions of both T1w reference and the T1w template. The following template was selected for spatial normalization: ICBM 152 Nonlinear Asymmetrical template version 2009c (Fonov, 2009) [RRID:SCR_008796; TemplateFlow ID: MNI152NLin2009cAsym].

Removed in the revised version: Functional data preprocessing

Earlier version: First, a reference volume and its skull-stripped version were generated using a custom methodology of fMRIPrep. Head-motion parameters with respect to the BOLD reference (transformation matrices, and six corresponding rotation and translation parameters) are estimated before any spatiotemporal filtering using mcflirt (Jenkinson et al., 2002) [FSL 6.0.5.1:57b01774]. BOLD runs were slice-time corrected to 0.96s (0.5 of slice acquisition range 0s-1.92s) using 3dTshift from AFNI (Cox & Hyde, 1997). The BOLD time-series (including slice-timing correction when applied) were resampled onto their original, native space by applying the transforms to correct for head-motion. These resampled BOLD time-series will be referred to as preprocessed BOLD in original space, or just preprocessed BOLD. The BOLD reference was then co-registered to the T1w reference using mri_coreg (FreeSurfer) followed by flirt (Jenkinson & Smith, 2001) [FSL 6.0.5.1:57b01774] with the boundary-based registration (Greve & Fischl, 2009) cost-function. Co-registration was configured with six degrees of freedom. Several confounding time-series were calculated based on the preprocessed BOLD: framewise displacement (FD), DVARS and three region-wise global signals. FD was computed using two formulations following Power (absolute sum of relative motions) (Power et al., 2014) and Jenkinson (relative root mean square displacement between affines, Jenkinson et al. (2002)). FD and DVARS are calculated for each functional run, both using their implementations in Nipype (following the definitions by Power et al. (2014)). The three global signals are extracted within the CSF, the WM, and the whole-brain masks. Additionally, a set of physiological regressors were extracted to allow for component-based noise correction (CompCor) (Behzadi et al., 2007). Principal components are estimated after high-pass filtering the preprocessed BOLD time-series (using a discrete cosine filter with 128s cut-off) for the two CompCor variants: temporal (tCompCor) and anatomical (aCompCor). tCompCor components are then calculated from the top 2% variable voxels within the brain mask. For aCompCor, three probabilistic masks (CSF, WM and combined CSF+WM) are generated in anatomical space. The implementation differs from that of Behzadi et al. (2007) in that instead of eroding the masks by 2 pixels on BOLD space, the aCompCor masks are subtracted a mask of pixels that likely contain a volume fraction of GM. This mask is obtained by thresholding the corresponding partial volume map at 0.05, and it ensures components are not extracted from voxels containing a minimal fraction of GM. Finally, these masks are resampled into BOLD space and binarized by thresholding at 0.99 (as in the original implementation). Components are also calculated separately within the WM and CSF masks. For each CompCor decomposition, the k components with the largest singular values are retained, such that the retained components’ time series are sufficient to explain 50 percent of variance across the nuisance mask (CSF, WM, combined, or temporal). The remaining components are dropped from consideration. Frames that exceeded a threshold of 0.5 mm FD or 1.5 standardised DVARS were annotated as motion outliers.

Revised version: Functional data preprocessing. First, a reference volume and its skull-stripped version were generated using a custom methodology of fMRIPrep. Head-motion parameters with respect to the BOLD reference (transformation matrices, and six corresponding rotation and translation parameters) are estimated before any spatiotemporal filtering using mcflirt (Jenkinson et al., 2002) [FSL 6.0.5.1:57b01774]. BOLD runs were slice-time corrected to 0.96s (0.5 of slice acquisition range 0s-1.92s) using 3dTshift from AFNI (Cox & Hyde, 1997). The BOLD time-series (including slice-timing correction when applied) were resampled onto their original, native space by applying the transforms to correct for head-motion. These resampled BOLD time-series will be referred to as preprocessed BOLD in original space, or just preprocessed BOLD. The BOLD reference was then co-registered to the T1w reference using mri_coreg (FreeSurfer) followed by flirt (Jenkinson & Smith, 2001) [FSL 6.0.5.1:57b01774] with the boundary-based registration (Greve & Fischl, 2009) cost-function. Co-registration was configured with six degrees of freedom. Several confounding time-series were calculated based on the preprocessed BOLD: framewise displacement (FD), DVARS and three region-wise global signals. FD was computed using two formulations following Power (absolute sum of relative motions) (Power et al., 2014) and Jenkinson (relative root mean square displacement between affines, Jenkinson et al. (2002)). FD and DVARS are calculated for each functional run, both using their implementations in Nipype (following the definitions by Power et al. (2014)). The three global signals are extracted within the CSF, the WM, and the whole-brain masks. Additionally, a set of physiological regressors were extracted to allow for component-based noise correction (CompCor) (Behzadi et al., 2007). Principal components are estimated after high-pass filtering the preprocessed BOLD time-series (using a discrete cosine filter with 128s cut-off) for the two CompCor variants: temporal (tCompCor) and anatomical (aCompCor). tCompCor components are then calculated from the top 2% variable voxels within the brain mask. For aCompCor, three probabilistic masks (CSF, WM and combined CSF+WM) are generated in anatomical space. The implementation differs from that of Behzadi et al. (2007) in that instead of eroding the masks by 2 pixels on BOLD space, the aCompCor masks are subtracted a mask of pixels that likely contain a volume fraction of GM. This mask is obtained by thresholding the corresponding partial volume map at 0.05, and it ensures components are not extracted from voxels containing a minimal fraction of GM. Finally, these masks are resampled into BOLD space and binarized by thresholding at 0.99 (as in the original implementation). Components are also calculated separately within the WM and CSF masks. For each CompCor decomposition, the k components with the largest singular values are retained, such that the retained components’ time series are sufficient to explain 50 percent of variance across the nuisance mask (CSF, WM, combined, or temporal). The remaining components are dropped from consideration. Frames that exceeded a threshold of 0.5 mm FD or 1.5 standardised DVARS were annotated as motion outliers.

The BOLD time-series were resampled into standard space, generating a preprocessed BOLD run in MNI152NLin2009cAsym space. First, a reference volume and its skull-stripped version were generated using a custom methodology of fMRIPrep. All resamplings can be performed with a single interpolation step by composing all the pertinent transformations (i.e., head-motion transform matrices, susceptibility distortion correction when available, and co-registrations to anatomical and output spaces). Gridded (volumetric) resamplings were performed using antsApplyTransforms (ANTs), configured with Lanczos interpolation to minimize the smoothing effects of other kernels (Lanczos, 1964). Non-gridded (surface) resamplings were performed using mri_vol2surf (FreeSurfer).

EEG sleep rhythms analysis

In line with other studies examining the coordination between SOs and spindles, particularly their spatial distribution (Massimini et al., 2004; Molle et al., 2011), EEG data from the F3 electrode were used to quantify the occurrence of SOs and spindles. Using the established detection algorithms (detailed below), we identified SOs and sleep spindles for each subject.

Added in the revised version: It is worth noting that the primary aim of EEG rhythm detection was to identify reliable event times for EEG-informed fMRI modelling. Detection was performed on the F3 electrode because this channel provided stable signal quality during prolonged nocturnal EEG-fMRI recordings. In our MR-compatible EEG setup, FCz was used as the online reference. Central electrodes close to this reference, and electrodes near the vertex that were in prolonged contact with the head coil during supine sleep, were more susceptible to reduced signal contrast, impedance drift, and pressure-related degradation of electrode-scalp contact. We therefore used F3 as a pragmatic choice to maximize reliable event timing in N2/3 sleep. This choice was not intended to characterize the full scalp topography of SOs or spindles, and it may underrepresent the centro-parietal fast spindle component. As a sensitivity analysis, we repeated the main EEG-informed fMRI GLM using Fz, a midline frontal electrode, with the same detection and modelling procedure.

Removed in the revised version: Detection of SOs

Earlier version: Data were first bandpass-filtered between 0.16 and 1.25 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). After identifying all positive-to-negative zero crossings, potential SOs were defined based on the interval between consecutive zero crossings (ranging from 0.8 s to 3 s). For each potential SO, we calculated the amplitude range (peak minus trough). Only those exceeding the 75th percentile in amplitude were labeled as SOs, following previous work (Schreiner et al., 2021).

Revised version: Detection of SOs. Data were first bandpass-filtered between 0.16 and 1.25 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). After identifying all positive-to-negative zero crossings, potential SOs were defined based on the interval between consecutive zero crossings, ranging from 0.8 s to 3 s. For each potential SO, we calculated the amplitude range as the peak minus the trough.

Added in the revised version: For each participant, the amplitude threshold was defined as the 75th percentile of candidate SO amplitude ranges observed during N2/3 sleep. This fixed N2/3-derived threshold was then applied unchanged across the recording for descriptive stage-wise summaries. Detected events were assigned to N1, N2/3 or REM according to the sleep-stage label at the event time. Only candidates exceeding this threshold were labelled as SOs, following previous work (Schreiner et al., 2021).

Removed in the revised version: Detection of sleep spindles

Earlier version: Data were bandpass-filtered between 12 and 16 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). The root mean square (RMS) of the filtered signal was computed with a 200 ms sliding time window. The amplitude threshold for spindles was set at the 75th percentile of RMS values. Any RMS values that surpassed this threshold for 0.5 s to 3 s were identified as spindles (Staresina et al., 2015).

Revised version: Detection of sleep spindles. Detection of sleep spindles. Data were bandpass-filtered between 12 and 16 Hz (Butterworth filter, order 3, bidirectional filtering for zero phase). The root mean square (RMS) of the filtered signal was computed with a 200 ms sliding time window. For each participant, the spindle threshold was defined as the 75th percentile of RMS values observed during N2/3 sleep. This fixed N2/3-derived threshold was then applied unchanged across the recording for descriptive stage-wise summaries. Detected events were assigned to N1, N2/3 or REM according to the sleep-stage label at the event time. RMS segments exceeding this threshold for 0.5 s to 3 s were identified as spindles (Staresina et al., 2015).

Removed in the revised version: Detection of SO-spindle couplings

Earlier version: From the detected SOs and spindles, we identified the peak time of each spindle. Within each SO interval, we checked for a spindle peak; if found, it was recorded as an SO-spindle coupling event. For every SO-spindle coupling event, an epoch was created time-locked to the SO trough (as the central reference), following Schreiner et al. (2021). We extracted data in a [-4 s to 4 s] window around this point, forming the epoch for each coupling event.

Revised version: Detection of SO-spindle couplings. From the detected SOs and spindles, we identified the peak time of each spindle. Within each SO interval, we checked whether a spindle peak occurred; if so, that SO was labelled as an SO-spindle coupling event. For descriptive stage-wise summaries, coupling events were assigned to the sleep stage of the corresponding SO trough. For every SO-spindle coupling event, an epoch was created time-locked to the SO trough as the central reference, following Schreiner et al. (2021). We extracted data in a [−4 s to 4 s] window around this point, forming the epoch for each coupling event. For the EEG-informed fMRI analyses, only SO, spindle and SO-spindle coupling events detected during N2/3 sleep were used.

Added in the revised version: The detection procedures described above were developed primarily for N2 and N3 sleep, where SOs, spindles and their coupling are physiologically expected and most reliably observed (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). Because percentile-based thresholds can otherwise force the detector to label events in every sleep stage, we did not estimate separate thresholds within N1 or REM. Instead, for each participant, all SO and spindle thresholds were defined from N2/3 sleep and then applied uniformly across the recording. Tables S1 and S3 report detailed statistical information on sleep rhythm and N2/3 events detection. The N1 and REM events detection reported in Tables S2 and S4, and illustrated in Fig. S2-S4, should therefore be interpreted as descriptive detector outputs under this fixed N2/3-derived criterion, rather than as evidence for canonical N2/3 SOs, spindles or physiological SO-spindle complexes in those stages. These detections were not used in the EEG-informed fMRI GLM or PPI analyses, which were restricted to N2/3 sleep.

Removed in the revised version: We note that the above methods for detecting SOs, spindles, and their couplings were originally developed for N2 and N3 sleep data, based on the specific characteristics of these stages. These methods are widely recognized in sleep research (Hahn et al., 2020; Helfrich et al., 2019; Helfrich et al., 2018; Ngo, Fell, & Staresina, 2020; Schreiner et al., 2022; Schreiner et al., 2021; Staresina et al., 2015; Staresina et al., 2023). However, because this percentile-based detection approach will inherently identify a certain number of events if applied to other stages (e.g., N1 and REM), the nature of these events in those stages remains unclear compared to N2/N3. We nevertheless identified and reported the detailed descriptive statistics of these sleep rhythms in all sleep stages, under the same operational definitions, both for completeness and as a sanity check. Within the same subject, there should be more SOs, spindles, and their couplings in N2/N3 than in N1 or REM (see also Figure S2-S4, Table S1-S4).

Removed in the revised version: Event-related potentials (ERP) analysis

Earlier version: After completing the detection of each sleep rhythm event, we performed ERP analyses for SOs, spindles, and coupling events in different sleep stages. Specifically, for SO events, we took the trough of the DOWN-state of each SO as the zero-time point, then extracted data in a [-2 s to 2 s] window from the broadband (0.1–30 Hz) EEG and used [-2 s to -0.5 s] for baseline correction; the results were then averaged across 107 subjects (see Fig. S2a). For spindle events, we used the peak of each spindle as the zero-time point and applied the same data extraction window and baseline correction before averaging across 107 subjects (see Fig. S2b). Finally, for SO-spindle coupling events, we followed the same procedure used for SO events (see Fig. 2a, Figs. S3–S4).

Revised version: Event-related potentials (ERP) analysis. After completing the detection of each sleep rhythm event, we performed ERP analyses for SOs, spindles, and coupling events in different sleep stages. Specifically, for SO events, we took the trough of the DOWN-state of each SO as the zero-time point, then extracted data in a [-2 s to 2 s] window from the broadband (0.1–30 Hz) EEG and used [−2 s to −0.5 s] for baseline correction; the results were then averaged across 107 subjects (see Fig. S2a). For spindle events, we used the peak of each spindle as the zero-time point and applied the same data extraction window and baseline correction before averaging across 107 subjects (see Fig. S2b). Finally, for SO-spindle coupling events, we followed the same procedure used for SO events (see Fig. 2a, Figs. S3–S4).

Removed in the revised version: Time-frequency analysis

Earlier version: We performed this analysis on the SO-spindle coupling epochs using the FieldTrip toolbox (Oostenveld et al., 2011). Fourier analysis was executed on the data, utilizing a sliding time window with a forward shift of 50 ms. The window length was defined as five cycles of the present frequency (frequency range: 1-30 Hz, step size: 1 Hz). Before Fourier analysis, windowed data segments underwent multiplication by a Hanning taper. Subsequently, power values were normalized to z-scores within the time range of [-4 s to 4 s], selecting a longer duration to prevent edge artifacts from low-frequency activity in the target time window of [-1.5 s to 1.5 s].

Revised version: Time-frequency analysis. We performed this analysis on the SO-spindle coupling epochs using the FieldTrip toolbox (Oostenveld et al., 2011). Fourier analysis was executed on the data, utilizing a sliding time window with a forward shift of 50 ms. The window length was defined as five cycles of the present frequency (frequency range: 1-30 Hz, step size: 1 Hz). Before Fourier analysis, windowed data segments underwent multiplication by a Hanning taper. Subsequently, power values were normalized to z-scores within the time range of [-4 s to 4 s], selecting a longer duration to prevent edge artifacts from low-frequency activity in the target time window of [-1.5 s to 1.5 s].

Removed in the revised version: Preferred phase analysis

Earlier version: All SO-spindle coupling epochs were filtered to the SO frequency band (0.16-1.25 Hz, 3rd-order Butterworth filter), and the instantaneous SO phase was determined through the Hilbert transform. Using the saved spindle peak times for each SO-spindle coupling, the preferred phase was deduced. The distribution of these phases across participants was then tested for uniformity using the Rayleigh test.

Revised version: Preferred phase analysis. All SO-spindle coupling epochs were filtered to the SO frequency band (0.16-1.25 Hz, 3rd-order Butterworth filter), and the instantaneous SO phase was determined through the Hilbert transform. Using the saved spindle peak times for each SO-spindle coupling, the preferred phase was deduced. The distribution of these phases across participants was then tested for uniformity using the Rayleigh test.

fMRI GLM analysis

All fMRI GLM analyses were conducted using only data from the N2/3 stages of NREM sleep. Segments of sleep data from these stages were concatenated for each participant. After pre-processing, the fMRI data was smoothed with an 8 mm FWHM kernel. All images underwent a temporal high-pass filter (width of 100 s), and the autocorrelation of the hemodynamic response was modelled via an AR(1) model. We then conducted General Linear Model (GLM) analyses to identify whole-brain activations associated with sleep rhythms and their coupling. We incorporated nuisance regressors estimated during the preprocessing with fMRIprep: six rigid-body motion-correction parameters identified during realignment (comprising three translation and rotation parameters), the mean White Matter, and the mean Cerebral Spinal Fluid signals. The cosine drift model was employed to detect and eliminate data fluctuations, thereby enhancing the accuracy of actual experimental effects. The GLMs were designed to evaluate the influence of experimental conditions on BOLD responses. All whole-brain analyses were thresholded using whole-brain FWE corrected at the cluster level (p < 0.05), with a cluster-forming voxel threshold of punc. < 0.001, unless otherwise specified.

The fMRI GLM analysis was designed to look for brain activations pertinent to SO, spindle, and SO-spindle coupling. Drawing from the EEG results concerning the timing of SO and spindle, we identified the trough of the SO as the onset for SO events and the peak of the spindle as the onset for spindle events. We chose the trough of the SO and the peak of the spindle as the onset points, primarily due to the larger amplitudes observed at these moments (thereby more robust to noise). This was also the common choice in previous EEG studies on SO and spindles (Ngo, Fell, & Staresina, 2020; Staresina et al., 2015; Staresina et al., 2023). To model SO-spindle coupling in fMRI, we followed Schreiner et al. (2021), and defining SO-spindle complex as the peak of spindle appeared within 1.5 seconds after SO trough. When such a peak was found, that particular SO was labelled as an SO-spindle event, with its SO trough labelled as the event onset. After convolving these events with the canonical Hemodynamic Response Function (HRF) of SPM and downsampling them to align with the fMRI time resolution, they were integrated as regressors in the design matrix for GLM analysis, allowing us to determine the effects of SO and spindle, as well as the SO-spindle coupling (Fig. 3b-d).

We have also conducted a control analysis aimed to distinguish brain activations associated with DOWN- vs. UP state of SOs. Drawing again from EEG data on SO and spindle, we randomly selected half of the SO events using the peak as the onset and the other half using the trough, to reduce collinearity between the two regressors. As with the main analysis, we identified the spindle peak as the onset for spindle events (Fig. S5).

ROI analysis

The objective of our ROI analysis was to test the activation of specific brain regions during distinct sleep rhythms. Within each ROI, the beta values at the subject-level were averaged across all voxels to facilitate subsequent statistical inferences (one-sample t test). We defined the mDMN and mPFC ROIs functionally based on the meta-analysis results in Neurosynth (Yarkoni et al., 2011), by the keyword “medial default mode network” and “ventromedial prefrontal cortex” respectively. In addition to mDMN and mPFC, we defined other ROIs anatomically, using the high-resolution Harvard-Oxford Atlas probabilistic atlas (Desikan et al., 2006).

Open-ended Cognitive State Functional Decoding

For functional decoding, we followed the approach in Margulies et al. (2016). We used the NeuroSynth (Yarkoni et al., 2011) meta-analytic database (www.neurosynth.org) to assess topic terms associated with the sleep rhythms activation patterns. ROI masks were created from the whole-brain activation patterns associated with the three sleep rhythms in Fig. 3b-d. These activation patterns were binarized using a threshold of p < 0.001, forming a brain mask for each sleep rhythm.

For each ROI mask, the analysis output was a z-statistic associated with a feature term, and the terms were ranked based on the weighted sum of their z-scores. The feature terms were derived from the most general set of 50 topic terms available in the NeuroSynth v3 dataset (https://github.com/NeuroanatomyAndConnectivity/gradient_analysis/blob/master/gradient_data/neurosynth/v3-topics-50-keys.txt).

From the 50 topics, 16 exceeded the threshold of z > 1.96 (p < 0.05). Two terms were removed as they did not represent coherent cognitive functions, leaving 14 relevant topic terms (Fig. S7). These included: episodic memory, declarative memory, working memory, task representation, language, learning, faces, visuospatial processing, category recognition, cognitive control, reading, cued attention, inhibition, and action.

PPI analysis

PPI (psychophysiological interaction) analysis is designed to quantify context-dependent alterations in effective connectivity between brain regions. For instance, it can determine whether hippocampus-based functional connectivity with the thalamus varies based on the occurrence of an SO-spindle coupling event. We conducted nine whole-brain PPI analyses with nilearn to investigate interactions based in the hippocampus, thalamus or mPFC during SO-spindle coupling, isolated SO and isolated spindle, utilizing the same ROIs. Both the main effects of physiological and physiological condition, as well as their interaction were modelled in the design matrix for PPI analysis. The main effect of physiological term produced functional connectivity maps of the entire brain in relation to the specified ROI (Fig. 4b-d). While the interaction term asked for the change of functional connectivity along with the SO-spindle coupling (Fig. 4e).